Pesticide-induced resurgence in brown planthopper is mediated by action on a suite of genes that promote juvenile hormone biosynthesis and female fecundity

Author response:

The following is the authors’ response to the original reviews.

eLife assessment

This useful manuscript reports mechanisms behind the increase in fecundity in response to sub-lethal doses of pesticides in the crop pest, the brown plant hopper. The authors hypothesize that the pesticide works by inducing the JH titer, which through the JH signaling pathway induces egg development. Evidence for this is, however, inadequate.

We greatly appreciate your valuable comments and constructive suggestions for our work. All in all, the manuscript has been carefully edited and improved following your suggestions. We also provide more evidence to support our statements by conducting new experiments. First, we found that also EB treatment of adult females can stimulate egg-laying. Second, EB treatment in female adults increases the number of mature eggs in the ovary and ovarioles. Third, EB treatment in females enhances the expression of the kr-h1 gene in the whole body of BPH. Finally, EB treatment in female adults increases the JHIII titer, but has no impact on the 20E titer.

Public Reviews:

Reviewer #1 (Public Review):

Summary:

Gao et al. have demonstrated that the pesticide emamectin benzoate (EB) treatment of brown planthopper (BPH) leads to increased egg-laying in the insect, which is a common agricultural pest. The authors hypothesize that EB upregulates JH titer resulting in increased fecundity.

Strengths:

The finding that a class of pesticide increases the fecundity of brown planthopper is interesting.

We greatly appreciate your positive comments on our work.

Weaknesses:

(1) EB is an allosteric modulator of GluCl. That means EB physically interacts with GluCl initiating a structural change in the cannel protein. Yet the authors' central hypothesis here is about how EB can upregulate the mRNA of GluCl. I do not know whether there is any evidence that an allosteric modulator can function as a transcriptional activator for the same receptor protein. The basic premise of the paper sounds counterintuitive. This is a structural problem and should be addressed by the authors by giving sufficient evidence about such demonstrated mechanisms before.

Thank you for your question. As the reviewer points out, EB physically interacts with its target protein GluCl and thus affects its downstream signaling pathway. In the manuscript, we reported that EB-treated brown planthoppers display increased expression of GluCl in the adult stage (Fig. 5A). Actually, there are many studies showing that insects treated with insecticides can increase the expression of target genes. For example, the relative expression level of the ryanodine receptor gene of the rice stem borer, Chilo suppressalis was increased 10-fold after treatment with chlorantraniliprole, an insecticide which targets the ryanodine receptor (Peng et al., 2017). Besides this, in Drosophila, starvation (and low insulin) elevates the transcription level of the sNPF and tachykinin receptors (Ko et al., 2015; Root et al., 2011). In brown planthoppers, reduction in mRNA and protein expression of a nicotinic acetylcholine receptor α8 subunit is associated with resistance to imidacloprid (Zhang et al., 2015). RNA interference knockdown of α8 gene decreased the sensitivity of N. lugens to imidacloprid (Zhang et al., 2015). Hence, expression of receptor genes can be regulated by diverse factors including insecticide treatment. In our case, we found that EB can upregulate its target gene GluCl. However, we did not claim that EB functions as transcriptional activator for GluCl, and we still do not know why EB treatment changes the expression of GluCl in the brown planthopper. Considering our experiments are lasting several days, it might be an indirect (or secondary) effect caused by other factors, which change the expression of GluCl gene upon EB action of the channel. One reason is maybe that the allosteric interaction with GluCl by EB makes it dysfunctional and the cellular response is to upregulate the channel/receptor to compensate. We have inserted text on lines 738 - 757 to explain these possibilities.

(2) I am surprised to see a 4th instar larval application or treatment with EB results in the upregulation of JH in the adult stages. Complicating the results further is the observation that a 4th instar EB application results in an immediate decrease in JH titer. There is a high possibility that this late JH titer increase is an indirect effect.

Thank you for your question. Treatment with low doses or sublethal doses of insecticides might have a strong and complex impact on insects (Gandara et al., 2024; Gong et al., 2022; Li et al., 2023; Martelli et al., 2022). We kept the 4th instar of brown planthoppers feeding on EB for four days. They will develop to 5th instar after four days treatment, which is the final nymphal stage of BPH. Since the brown planthopper is a hemimetabolous insect, we cannot rule out the possibility that an indirect effect of treatment with EB results in the upregulation of JH in the adult stages. In this new revised manuscript, we investigated the impact of EB treatment in the adult stage. We found that female adults treated with EB also laid more eggs than controls (Figure 1-figure supplement 1A). The following experiments were performed in adults to address how EB treated stimulates egg-laying in adult brown planthopper.

(1) We found that EB treatment in adults increases the number of mature eggs in ovary (new Figure 2-figure supplement 1). We add this results in lines 234 – 238 and 281-285.

(2) We measured the JH titer after the female adults had been treated with EB. We found that EB can also increase the JH titer but has no impact on the 20E titer in the female adult (Figure 3-S3A and B). We add this results in lines 351 – 356 and 281-285.

(3) EB treatment in adults increases the gene expression of JHAMT and Kr-h1 (Figure 3-S3C and D). We add this results in lines 378 – 379, lines 387-390 and lines 457-462.

(3) The writing quality of the paper needs improvement. Particularly with respect to describing processes and abbreviations. In several instances the authors have not adequately described the processes they have introduced, thus confusing readers.

Thank you for your suggestion. We have thoroughly revised the paper to improve clarity.

(4) In the section 'EB promotes ovarian development' the authors have shown that EB treatment results in increased detention of eggs which contradicts their own results which show that EB promotes egg laying. Again, this is a serious contradiction that nullifies their hypothesis.

Thank you for pointing this out. We revised the figure 2B to show number of mature eggs in the ovary. The number of mature eggs in ovaries of females that fed on EB was higher than in control females. We also show that BPH fed with EB laid more eggs than controls. Thus, our results suggest that EB promotes ovary maturation (and egg production) and also increases egg laying (Figure 1 and Table S1). Thus, we found that EB treatment can increase both the production of eggs and increase egg laying. We add this results in lines 234 – 238.

(5) Furthermore, the results suggest that oogenesis is not affected by EB application. The authors should devote a section to discussing how they are observing increased egg numbers in EB-treated insects while not impacting Oogenesis.

Thank you for your suggestions, and apologies for the lack of clarity in our initial explanation. First, we found that EB treatment led to an increase in the number of eggs laid by female brown planthoppers (Figure 1). Through dissection experiments, we observed that EB-treated females had more mature eggs in their ovaries (Figure 2A and B), indicating that the increased egg-laying was due to a larger production of mature eggs in the ovaries after EB treatment. This is now explained on lines 229-238.

Additionally, since there is no systematic description of oogenesis in the brown planthopper, we were the first to observe the oogenesis process in this species using immunohistochemistry and laser confocal microscopy. Based on the developmental characteristics, we defined the different stages of oogenesis (Figure 2C, Figure 2-figure supplement 2). We did not observe any significant effect of EB treatment on the various stages of oogenesis, indicating that EB treatment does not impair normal egg development (Figure 2D). Instead, the increase in vitellogenin accelerates the production of mature eggs. This is now explained on lines 243-262.

During the maturation process, eggs require uptake of vitellogenin, and an increase in vitellogenin (Vg) content can accelerate egg maturation, producing more mature eggs. Our molecular data suggest that EB treatment leads to an upregulation of vg expression. Based on these findings, we conclude that the increase in egg-laying caused by EB treatment is due to the upregulation of vg (Figure 3I), which raises vitellogenin content, promoting the uptake of vitellogenin by maturing eggs and resulting in the production of more mature eggs. We have revised the text on lines 389-395 to clarify this point.

(6) Met is the receptor of JH and to my understanding, remains mostly constant in terms of its mRNA or protein levels throughout various developmental periods in many different insects. Therefore, the presence of JH becomes the major driving factor for physiological events and not the presence of the receptor Met. Here the authors have demonstrated an increase in Met mRNA as a result of EB treatment. Their central hypothesis is that EB increases JH titer to result in enhanced fecundity. JH action will not result in the activation of Met. Although not contradictory to the hypothesis, the increase in mRNA content of Met is contrary to the findings of the JH field thus far.

Thank you for your comment. Our results showed that EB treatment can mildly increase (about 2-fold) expression of the Met gene in brown planthoppers (Figure 3G). And our data indicated that Met and FAMeT expression levels were not influenced so much by EB compared with kr-h1 and vg (Figure 3H and I). We agree that JH action will not result in the increase of Met. However, we cannot rule out the possibility of other factors (indirect effects), induced by EB treatment that increase the mRNA expression level of Met. One recent paper reported that downregulation of transcription factor CncC will increase met expression in beetles (see Figure 6A in this reference) (Jiang et al., 2023). Many studies have reported that insecticide treatment will activate the CncC gene signaling pathway, which regulates detoxification gene expression (Amezian et al., 2023; Fu et al., 2024; Hu et al., 2021). Hence, it is possible that EB might influence the CncC gene pathway which then induces met expression. This EB effect on met upregulation may be similar to the upregulation of GluCl and some other secondary effects. We have discussed this on lines 725-738.

(7) As pointed out before, it is hard to rationalize how a 4th instar exposure to EB can result in the upregulation of key genes involved in JH synthesis at the adult stage. The authors must consider providing a plausible explanation and discussion in this regard.

Thank you for your comments. It must be mentioned that although we exposed the BPH to EB at 4th instar, we make the insect feed on the EB-treated rice plants for four days. After that, the insect will develop into 5th instar, the final nymphal stage of brown planthopper. Since brown planthoppers do not have a pupal stage, this might cause the EB presented to the insects last a longer time even in the adult stage. Besides this, we found that EB treatment will increase the weight of adult females (Figure 1-figure supplement 3E and F), which indicates that EB might increase food intake in BPHs that might produce more insulin peptide. Insulin might increase the JH synthesis at the adult stage. In our revised study we also investigate EB impairment in adult BPHs. We found that, similar to the nymphal stage, EB treatment in adult BPHs also increases the egg laying. Furthermore, the JH titer was increased after treatment of BPH with EB in adults. Besides this, GluCl and kr-h1 genes were also up-regulated after EB treatment in the adult stage. We have discussed this on lines 739-746.

(8) I have strong reservations against such an irrational hypothesis that Met (the receptor for JH) and JH-Met target gene Kr-h1 regulate JH titer (Line 311, Fig 3 supplemental 2D). This would be the first report of such an event on the JH field and therefore must be analysed in depth. I strongly suggest the authors remove such claims from the manuscript without substantiating it.

Thank you for your suggestions and comments. We have changed our claims in this revised MS. We found that EB treatment can enhance Kr-h1 expression. We have no evidence to support that JH can induce met expression. We have rewritten the manuscript to avoid confusion (see text on lines 725-735).

(9) Kr-h1 is JH/Met target gene. The authors demonstrate that silencing of Kr-h1 results in inhibition of FAMeT, which is a gene involved in JH synthesis. A feedback loop in JH synthesis is unreported. It is the view of this reviewer that the authors must go ahead with a mechanistic detail of Kr-h1 mediated JH upregulation before this can be concluded. Mere qPCR experiments are not sufficient to substantiate a claim that is completely contrary to the current understanding of the JH signalling pathway.

Thank you for your suggestions and comments. We agree that only qPCR experiments are not enough to provide this kind of claim. More evidences need to be provided to support this. We have revised the MS to avoid confusion (see text on lines 725-735).

(10) The authors have performed knockdowns of JHAMT, Met, and Kr-h1 to demonstrate the effect of these factors on fecundity in BPH. Additionally, they have performed rescue experiments with EB application on these knockdown insects (Figure 3K-M). This, I believe, is a very flawed experiment. The authors demonstrate EB works through JHAMT in upregulating JH titer. In the absence of JHAMT, EB application is not expected to rescue the phenotype. But the authors have reported a complete rescue here. In the absence of Met, the receptor of JH, either EB or JH is not expected to rescue the phenotype. But a complete rescue has been reported. These two experimental results contradict their own hypothesis.

Thank you for your comments. We thought that this rescue is possible since knockdown of the genes is incomplete when using dsRNA injection (and residual gene expression allows for EB action). It is not a total knockout and actually, these genes still have a low level of expression in the dsRNA-injected insects. Since EB can upregulate the expression of JHAMT, Met, and Kr-h1, it is reasonable that EB treatment can rescue the down-regulation effects of these three genes and make fecundity completely rescued. We have clarified this on lines 411-413).

(11) A significant section of the paper deals with how EB upregulates JH titer. JH is a hormone synthesized in the Corpora Allata. Yet the authors have chosen to use the whole body for all of their experiment. Changes in the whole body for mRNA of those enzymes involved in JH synthesis may not reflect the situation in Corpora Allata. Although working with Corpora Allata is challenging, discarding the abdomen and thorax region and working with the head and neck region of the insect is easily doable. Results from such sampling are always more convincing when it comes to JH synthesis studies.

Thank you for your suggestions. Because the head is very difficult to separate from the thorax region in brown planthoppers as you can see in Author response image 1. We are now trying to answer how EB regulates JH synthesis using Drosophila as a model.

The brown planthopper

(12) The phenomenon reported was specific to BPH and not found in other insects. This limits the implications of the study.

Thank you for your comments. The brown planthopper is a serious insect pest on rice in Asia. Our findings can guide the use of this insecticide in the field. Besides this, our findings indicated that EB, which targets GluCl can impair the JH titer. Our findings added new implications for how a neuronal system influences the JH signaling pathway. We will further investigate how EB influences JH in the future and will use Drosophila as a model to study the molecular mechanisms.

(13) Overall, the molecular experiments are very poorly designed and can at best be termed superficial. There are several contradictions within the paper and no discussion or explanation has been provided for that.

Thank you for your comments. We have revised the paper according to your suggestions and added further explanation of our results in the discussion parts and hope the conclusions are better supported in the new version. We have discussed this on lines 725-746 and 778-799.

Reviewer #2 (Public Review):

The brown plant hopper (BPH) is a notorious crop pest and pesticides are the most widespread means of controlling its population. This manuscript shows that in response to sublethal doses of the pesticide (EB), BPH females show enhanced fecundity. This is in keeping with field reports of population resurgence post-pesticide treatment. The authors work out the mechanism behind this increase in fecundity. They show that in response to EB exposure, the expression of its target receptor, GluCl, increases. This, they show, results in an increase in the expression of genes that regulate the synthesis of juvenile hormone (JH) and JH itself, which, in turn, results in enhanced egg-production and egg-laying. Interestingly, these effects of EB exposure are species-specific, as the authors report that other species of plant hoppers either don't show enhanced fecundity or show reduced fecundity. As the authors point out, it is unclear how an increase in GluCl levels could result in increased JH regulatory genes.

We greatly appreciate your valuable comments and constructive suggestion to our work. We will try to figure out how EB interacts with its molecular target GluCl and then increases JH regulatory genes in the future work using Drosophila as models.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

Overall, the molecular experiments are very poorly designed and can at best be termed superficial. There are several contradictions within the paper and no discussion or explanation has been provided for that.

The authors should consider a thorough revision.

Thank you for your comments. We have thoroughly revised the paper according to your suggestions and added further experiments and explanations of our results in the discussion parts.

Reviewer #2 (Recommendations For The Authors):

It would help the reader to have more schematics along with the figures. The final figure is helpful, but knowing the JH pathway, and where it acts would help with the interpretations as one reads the manuscript and the figures. The pathways represented in 4N or 5J are helpful but could be improved upon for better presentation.

It would be nice to have some discussion on how the authors think EB exposure results in an increase in GluCl expression, and how that in turn affects the expression of so many genes.

Thank you for your comments. We have thoroughly revised the paper according to your suggestions and added further experiments and explanations of how we think EB exposure results in an increase in JH titer and other genes in the discussion parts. We have added the test on lines 753-761.

References

Amezian, D., Fricaux, T., de Sousa, G., Maiwald, F., Huditz, H.-I., Nauen, R., Le Goff, G., 2023. Investigating the role of the ROS/CncC signaling pathway in the response to xenobiotics in Spodoptera frugiperda using Sf9 cells. Pesticide Biochemistry and Physiology 195, 105563.

Fu, B., Liang, J., Hu, J., Du, T., Tan, Q., He, C., Wei, X., Gong, P., Yang, J., Liu, S., Huang, M., Gui, L., Liu, K., Zhou, X., Nauen, R., Bass, C., Yang, X., Zhang, Y., 2024. GPCR–MAPK signaling pathways underpin fitness trade-offs in whitefly. Proceedings of the National Academy of Sciences 121, e2402407121.

Gandara, L., Jacoby, R., Laurent, F., Spatuzzi, M., Vlachopoulos, N., Borst, N.O., Ekmen, G., Potel, C.M., Garrido-Rodriguez, M., Böhmert, A.L., Misunou, N., Bartmanski, B.J., Li, X.C., Kutra, D., Hériché, J.-K., Tischer, C., Zimmermann-Kogadeeva, M., Ingham, V.A., Savitski, M.M., Masson, J.-B., Zimmermann, M., Crocker, J., 2024. Pervasive sublethal effects of agrochemicals on insects at environmentally relevant concentrations. Science 386, 446-453.

Gong, Y., Cheng, S., Desneux, N., Gao, X., Xiu, X., Wang, F., Hou, M., 2022. Transgenerational hormesis effects of nitenpyram on fitness and insecticide tolerance/resistance of Nilaparvata lugens. Journal of Pest Science.

Hu, B., Huang, H., Hu, S., Ren, M., Wei, Q., Tian, X., Esmail Abdalla Elzaki, M., Bass, C., Su, J., Reddy Palli, S., 2021. Changes in both trans- and cis-regulatory elements mediate insecticide resistance in a lepidopteron pest, Spodoptera exigua. PLOS Genetics 17, e1009403.

Jiang, H., Meng, X., Zhang, N., Ge, H., Wei, J., Qian, K., Zheng, Y., Park, Y., Reddy Palli, S., Wang, J., 2023. The pleiotropic AMPK–CncC signaling pathway regulates the trade-off between detoxification and reproduction. Proceedings of the National Academy of Sciences 120, e2214038120.

Ko, K.I., Root, C.M., Lindsay, S.A., Zaninovich, O.A., Shepherd, A.K., Wasserman, S.A., Kim, S.M., Wang, J.W., 2015. Starvation promotes concerted modulation of appetitive olfactory behavior via parallel neuromodulatory circuits. eLife 4, e08298.

Li, Z., Wang, Y., Qin, Q., Chen, L., Dang, X., Ma, Z., Zhou, Z., 2023. Imidacloprid disrupts larval molting regulation and nutrient energy metabolism, causing developmental delay in honey bee Apis mellifera. eLife

Martelli, F., Hernandes, N.H., Zuo, Z., Wang, J., Wong, C.-O., Karagas, N.E., Roessner, U., Rupasinghe, T., Robin, C., Venkatachalam, K., Perry, T., Batterham, P., Bellen, H.J., 2022. Low doses of the organic insecticide spinosad trigger lysosomal defects, elevated ROS, lipid dysregulation, and neurodegeneration in flies. eLife 11, e73812.

Peng, Y.C., Sheng, C.W., Casida, J.E., Zhao, C.Q., Han, Z.J., 2017. Ryanodine receptor genes of the rice stem borer, Chilo suppressalis: Molecular cloning, alternative splicing and expression profiling. Pestic. Biochem. Physiol. 135, 69-77.

Root, Cory M., Ko, Kang I., Jafari, A., Wang, Jing W., 2011. Presynaptic facilitation by neuropeptide signaling mediates odor-driven food search. Cell 145, 133-144.

Zhang, Y., Wang, X., Yang, B., Hu, Y., Huang, L., Bass, C., Liu, Z., 2015. Reduction in mRNA and protein expression of a nicotinic acetylcholine receptor α8 subunit is associated with resistance to imidacloprid in the brown planthopper, Nilaparvata lugens. Journal of Neurochemistry 135, 686-694.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Pesticide-induced resurgence, increases in pest insect populations following pesticide application, is a serious threat to the sustainable control of many highly damaging crop pests. Resurgence can result from pesticide-enhanced pest reproduction, however, the molecular mechanisms mediating this process remain unresolved. Here we show that brown planthopper (BPH) resurgence following emamectin benzoate (EB) exposure results from the coordinated action of a diverse suite of actors that regulate juvenile hormone (JH) levels, resulting in increased JH titer in adult females and enhanced fecundity. Following reports of BPH resurgence in rice crops when this species is exposed to EB, we demonstrate that EB treatment results in profound changes in female BPH fitness including enhanced ovarian development and elevated egg production. This enhanced reproductive fitness results from the EB-mediated upregulation of key genes involved in the regulation of JH, including JHAMT, Met and Kr-h1 and the downregulation of allatostatin (AstA) and allatostatin receptor (AstAR) expression. The remodulation of gene expression following EB exposure is dependent on the action of this insecticide on its molecular target the glutamate-gated chloride channel (GluCl) receptor. Collectively, these results provide mechanistic insights into the regulation of negative pesticide-induced responses in insects and reveal the key actors involved in the JH-signaling pathway that underpin pesticide resurgence.

Revised version: Pesticide-induced resurgence, increases in pest insect populations following pesticide application, is a serious threat to the sustainable control of many highly damaging crop pests. Resurgence can result from pesticide-enhanced pest reproduction, however, the molecular mechanisms mediating this process remain unresolved. Here we show that brown planthopper (BPH) resurgence in rice crops following exposure to the pesticide emamectin benzoate (EB) results from the coordinated action of a suite of genes that regulate juvenile hormone (JH) levels, resulting in increased JH titer in adult females and enhanced fecundity. We demonstrate that EB treatment at sublethal levels results in profound changes in female BPH fitness including increased egg maturation and oviposition. This enhanced reproductive fitness results from the EB-mediated upregulation of key genes involved in the regulation of JH, including JHAMT and Kr-h1 and the downregulation of allatostatin (AstA) and allatostatin receptor (AstAR) expression. AstA signaling is known to inhibit the production of JH in the corpora allata and hence EB exposure diminishes this inhibitory action. We find that the changes in gene expression following EB exposure are caused by the action of this insecticide on its molecular target, the glutamate-gated chloride channel (GluCl). Collectively, these results provide mechanistic insights into the regulation of negative pesticide-induced responses in insects and reveal the key actors involved in the JH-signaling pathway that underpin pesticide resurgence.

Introduction

Earlier version: Chemical pesticides remain the primary means of controlling many of the world’s most damaging arthropod crop pests (1, 2). However, pesticide applications can result in pest resurgence, increases in pest insect populations that exceed natural, untreated population sizes, following an initial reduction of the pest population (3–5). Two mechanisms have been implicated in pest resurgence - the loss of beneficial insects including natural enemies and pesticide-enhanced pest reproduction (3). In the case of the latter, several pesticides, such as the insecticides triazophos, deltamethrin and the fungicide jinggangmycin, have been reported to stimulate pest reproduction (6–8). Pesticide-enhanced pest reproduction has been linked to changes in physiology and biochemistry of pest organisms after exposure to pesticides (3, 4). However, the molecular mechanisms underlying enhanced reproduction associated with pest resurgence remain poorly resolved.

Revised version: Chemical pesticides remain the primary means of controlling many of the world’s most damaging arthropod crop pests (Janssen and van Rijn, 2021; Wu et al., 2018a). However, pesticide applications can result in pest resurgence, increases in pest insect populations that exceed natural, untreated population sizes, following an initial reduction of the pest population (Gandara et al., 2024; Guedes et al., 2016; Hardin et al., 1995; Wu et al., 2020). Two mechanisms have been implicated in pest resurgence - the loss of beneficial insects including natural enemies, and pesticide-enhanced reproduction in the pest insect (Wu et al., 2020). In the case of the latter, several pesticides, such as the insecticides triazophos, deltamethrin and the fungicide jinggangmycin, have been reported to stimulate pest reproduction (Ge et al., 2010; Wu et al., 2018b; Zhang et al., 2014). Pesticide-enhanced pest reproduction has been linked to changes in physiology and biochemistry of pest organisms after exposure to pesticides (Guedes et al., 2016; Wu et al., 2020). However, the molecular mechanisms underlying enhanced reproduction associated with pest resurgence remain poorly resolved.

Earlier version: The brown planthopper (BPH), Nilaparvata lugens (Stål), is a notorious pest of rice crops throughout Asia causing annual losses of ∼300 million dollars across major rice producing countries (2, 3). BPH inhibits the growth of rice plants by feeding, and also transmits highly damaging plant viruses including rice grassy stunt virus and rice ragged stunt virus (9). Currently, chemical insecticides play an indispensable role in the control of BPH due to their efficiency, rapid effect, and low cost. However, due to the widespread and intensive use of chemical insecticides, BPH has developed resistance to the majority of compounds used for control (2, 10).

Revised version: The brown planthopper (BPH), Nilaparvata lugens (Stål), is a notorious pest of rice crops throughout Asia causing annual losses of ∼300 million dollars across major rice producing countries (Wu et al., 2020; Wu et al., 2018a). BPHs inhibit the growth of rice plants by feeding, and also transmits highly damaging plant viruses including rice grassy stunt virus and rice ragged stunt virus (Sōgawa, 1982). Currently, chemical insecticides play an indispensable role in the control of BPH due to their efficiency, rapid effect, and low cost. However, due to the widespread and intensive use of chemical insecticides, BPH has developed resistance to the majority of compounds used for control (Wu et al., 2018a; Zeng et al., 2023).

Earlier version: Emamectin benzoate (EB) and abamectin are avermectin pesticides, and act as allosteric modulators of insect glutamate gated chloride channels (GluCls), inhibiting muscle contractions that lead to the cessation of insect feeding and subsequent death (11). These insecticides exhibit particularly strong activity against Lepidoptera such as the rice leaffolder, Cnaphalocrocis medinalis Guénee, an important foliage-feeding insect which attacks rice during the vegetative stage (12). Both BPH and the rice leaffolder are migratory pests with overlapping migratory paths, however, their occurrence period in the field differs by approximately one month, with leaffolders appearing earlier than BPH. Therefore, the use of EB to control rice leaffolder has the potential to impact BPH arriving later, via exposure to sublethal concentrations of this compound. In this regard, we have observed that when farmers use EB and abamectin to control leaffolders on rice crops in China, BPH outbreaks frequently occur in the same field. While sublethal doses of certain pesticides have been shown to enhance fecundity in BPH, including the insecticides triazophos and deltamethrin (6, 7, 13, 14) and the fungicides carbendazim and jinggangmycin (8), whether avermectins trigger resurgence in BPH via insecticide-enhanced reproduction remains unclear.

Revised version: Emamectin benzoate (EB) and abamectin are avermectin pesticides, and act as allosteric modulators of insect glutamate gated chloride channels (GluCls), inhibiting muscle contractions that lead to the cessation of insect feeding and subsequent death (Ishaaya et al., 2002). These insecticides exhibit particularly strong activity against Lepidoptera such as the rice leaffolder, Cnaphalocrocis medinalis Guénee, an important foliage-feeding insect which attacks rice during the vegetative stage (Chintalapati et al., 2016). Both BPH and the rice leaffolder are migratory pests with overlapping migratory paths, however, their occurrence period in the field differs by approximately one month, with leaffolders appearing earlier than BPH. Therefore, the use of EB to control rice leaffolder has the potential to impact BPHs arriving later, via exposure to sublethal concentrations of this compound. In this regard, we have observed that when farmers use EB and abamectin to control leaffolders on rice crops in China, BPH outbreaks frequently occur in the same field. While sublethal doses of certain pesticides have been shown to enhance fecundity in BPHs, including the insecticides triazophos and deltamethrin (Ge et al., 2013; Ge et al., 2010; Zhang et al., 2014; Zhang et al., 2022b) and the fungicides carbendazim and jinggangmycin (Wu et al., 2018b), whether avermectins trigger resurgence in BPH via insecticide-enhanced reproduction remains unclear.

Earlier version: Reproduction in insects is influenced by external factors such as light (15), temperature (16), humidity (17) and nutrition (18), and endogenous factors such as the juvenile hormone (JH) (19), ecdysone (20), insulin (21) and TOR (22) pathways (23, 24). Of these, JH has been particularly implicated in insecticide-induced enhanced fecundity, with triazophos and deltamethrin treatments leading to increased circulating JH III titers in BPH females (3). JH is synthesized and secreted by the corpora allata in insects (25), and can promote reproduction by regulating the synthesis and secretion of vitellogenin (vg) in the female fat body, and stimulating the absorption of vg by the developing oocyte (19). However, the regulation of JH is complex (19) and the key actors involved in JH-mediated pesticide-enhanced reproduction remains an open question. Previous finding reported that GluCl receptor is involved in JH biosynthesis in cockroach (26–28). However, we know little whether insecticide which target on the GluCl can stimulate reproduction in insect pest and how GluCl receptor regulates JH production is still unknown.

Revised version: Reproduction in insects is influenced by external factors such as light (Wang et al., 2021), temperature (Meiselman et al., 2022), humidity (Roy et al., 2015) and nutrition (Smykal and Raikhel, 2015), and endogenous factors such as the juvenile hormone (JH) (Santos et al., 2019), ecdysone (Hun et al., 2022), insulin (Ling and Raikhel, 2018) and target of rapamycin (TOR) pathways (Ahmed et al., 2020; Du et al., 2022; Lu et al., 2016). Of these, JH has been particularly implicated in insecticide-induced enhanced fecundity, with triazophos and deltamethrin treatments leading to increased circulating JH III titers in BPH females (Wu et al., 2020). JH is synthesized and secreted by the corpora allata in insects (Luo et al., 2021), and can promote reproduction by regulating the synthesis and secretion of vitellogenin (Vg) in the female fat body, and stimulating the absorption of Vg by the developing oocyte (Santos et al., 2019). However, the regulation of JH is complex (Santos et al., 2019) and the key actors involved in JH-mediated pesticide-enhanced reproduction remains an open question. Previous finding reported that the GluCl receptor is involved in JH biosynthesis in cockroach (Chiang et al., 2002a; Chiang et al., 2002b; Liu et al., 2005). However, it is still unknown whether insecticides that target the GluCl can increase reproduction in pest insects and how the GluCl receptor regulates JH production.

Earlier version: In this study, we used a diverse range of approaches to investigate the impact of sublethal doses of avermectins on BPH fecundity, and unravel the molecular mechanisms mediating enhanced reproduction following exposure to this insecticide class. We show that avermectin exposure results in profound changes in the expression of a key suite of genes that in combination regulate JH, resulting in increased JH titer in adult females, which promotes fecundity. Interestingly, we found that GluCl channel modulators can inhibit Ast/AstAR expression and then promotes JH increasing.

Revised version: In this study, we used a range of diverse approaches to investigate the impact of sublethal doses of avermectins on BPH fecundity, and unravel the molecular mechanisms mediating enhanced reproduction following exposure to this insecticide class. We show that avermectin exposure results in profound changes in the expression of a key suite of genes that in combination regulate JH, resulting in increased JH titer in adult females, which promotes fecundity. Interestingly, we found that GluCl channel modulators can inhibit AstA/AstAR expression and thereby promote JH production.

Results

GluCl allosteric modulators (emamectin benzoate and abamectin) stimulate fecundity of female N. lugens

Earlier version: To investigate whether GluCl modulators affect fecundity in BPH, we first determined the sub-lethal and median lethal doses of emamectin benzoate (EB) to 4th instar nymphs, newly emerged males and females of BPH (Table S1). For this we employed two different bioassay methods, the rice seedling dip bioassay method and topical application bioassay method (2, 29), in order to assess both the systemic and contact toxicity of these insecticides (Table S1). We then systemically treated 4th instar nymphs of BPH with the estimated LC15 and LC50 concentrations of EB and examined the fecundity of BPH after these individuals mated with treated or untreated individuals. We use the term “t” to represent individuals treated with EB and “ck” to indicate control individuals that were treated with insecticide diluent minus insecticide. After treatment with the LC15 and LC50 concentrations of EB, the number of eggs laid per female of BPH in ♀t × ♂t crosses increased by 1.48 and 1.40 times compared with control ♀ck × ♂ck crosses (Figure 1A); the number of eggs laid per female of BPH in ♀t × ♂ck crosses increased by 1.53 and 2.07 times compared with control crosses (Figure 1B); However, the number of eggs laid by per female of BPH in ♀ck × ♂t crosses did not increase significantly compared to control ♀ck × ♂ck crosses (Figure 1C).

Revised version: To investigate whether GluCl modulators affect fecundity in BPH, we first determined the sub-lethal and median lethal doses of emamectin benzoate (EB) on 4th instar nymphs, as well as on newly emerged males and females of BPH (Table S1). For this we employed two different bioassay methods, the rice seedling dip bioassay method and topical application bioassay method (Wu et al., 2018a; Yao et al., 2017), in order to assess both the systemic and contact toxicity of these insecticides (Table S1). We then systemically treated 4th instar nymphs of BPH, newly emerged males and females with the estimated LC15 or LC50 concentrations of EB and examined the fecundity of BPH after these individuals mated with treated or untreated individuals. We use the term “t” to represent individuals treated with EB and “ck” to indicate control individuals that were treated with insecticide diluent minus insecticide. First, we tested whether treatment of BPH with EB in nymphal stage impairs the fecundity of BPHs at the adult stage. For 4th instar nymphs of BPHs, after treatment with the LC15 and LC50 concentrations of EB, the number of eggs laid per female of BPHs in ♀t × ♂t crosses increased by 1.48 and 1.40 times compared with control ♀ck × ♂ck crosses (Figure 1A); the number of eggs laid per female of BPH in ♀t × ♂ck crosses increased by 1.53 and 2.07 times compared with control crosses (Figure 1B); However, the number of eggs laid by per female of BPH in ♀ck × ♂t crosses did not increase significantly compared to control ♀ck × ♂ck crosses (Figure 1C).

Earlier version: Fecundity of BPH following exposure to sub-lethal (LC15) and median lethal (LC50) concentrations of emamectin benzoate following system application bioassays (A: ♀t ×♂t; B: ♀t ×♂ck; C: ♀ck ×♂t) and topical application bioassays (D: ♀t ×♂t; E: ♀t ×♂ck; F: ♀ck ×♂t), respectively. The letter “t” represents treatment with insecticide, while “ck” indicates controls that was not treated with insecticide. All data are presented as the mean ± s.e.m. Different lower-case letters above the bars indicate significant differences (One-way ANOVA with Tukey’s Multiple Range Test, p < 0.05).

Revised version: Fecundity of BPH following exposure to sub-lethal (LC15) and median lethal (LC50) concentrations of emamectin benzoate following system application bioassays (A: ♀t ×♂t; B: ♀t ×♂ck; C: ♀ck ×♂t) and topical application bioassays (D: ♀t ×♂t; E: ♀t ×♂ck; F: ♀ck ×♂t), respectively. The letter “t” represents treatment with insecticide, while “ck” indicates controls that was not treated with insecticide. These were exposed as 4th instar nymphs. All data are presented as the mean ± s.e.m. Different lower-case letters above the bars indicate significant differences (One-way ANOVA with Tukey’s Multiple Range Test, p < 0.05).

Earlier version: Exposure of 4th instar nymphs with the LC15 and LC50 concentrations of EB in contact bioassays also significantly stimulated fecundity. After treatment with the LC15 and LC50 concentrations of EB, the number of eggs laid by per female of BPH in ♀t × ♂t crosses increased by 1.18 and 1.26 times compared with the control (♀ck × ♂ck) (Figure 1D); The number of eggs laid per female of BPH in ♀t × ♂ck crosses increased by 1.27 and 1.56 times compared with the control crosses (Figure 1E); However, there was no significant difference in egg-laying number between ♀ck × ♂t crosses and controls (♀ck × ♂ck) (Figure 1F). These results reveal that EB can stimulate the fecundity of females following both the systemic and contact routes of exposure.

Revised version: Exposure of 4th instar nymphs to the LC15 and LC50 concentrations of EB in contact bioassays also significantly stimulated fecundity. After treatment with the LC15 and LC50 concentrations of EB, the number of eggs laid by per female of BPH in ♀t × ♂t crosses increased by 1.18 and 1.26 times compared with the control (♀ck × ♂ck) (Figure 1D); The number of eggs laid per female of BPH in ♀t × ♂ck crosses increased by 1.27 and 1.56 times compared with the control crosses (Figure 1E); However, there was no significant difference in number of eggs laid between ♀ck × ♂t crosses and controls (♀ck × ♂ck) (Figure 1F). These results reveal that EB stimulates the fecundity of females following both the systemic and contact routes of exposure.

Earlier version: We next examined whether EB treatment of adult BPH also stimulates reproduction. Indeed, treating newly emerged adults with the LC15 and LC50 concentrations of EB significantly stimulated the number of eggs laid per female (Figure 1-figure supplement 1A). Furthermore, sub-lethal exposure of 4th instar BPH nymphs to another GluCl allosteric modulator, abamectin (LC15 and LC50 concentrations) was also found to significantly enhance reproduction (Figure 1-figure supplement 1B).

Revised version: Next, we examined whether EB treatment of adult BPHs also stimulates reproduction. Indeed, treating newly emerged adults with the LD15 and LD50 concentrations of EB significantly stimulated the number of eggs laid per female (Figure 1-figure supplement 1A). Furthermore, sub-lethal exposure of 4th instar BPH nymphs to another GluCl allosteric modulator, abamectin (LC15 and LC50 concentrations) was also found to significantly enhance reproduction (Figure 1-figure supplement 1B).

Added in the revised version: (A) Fecundity of BPH when newly emerged adults were treated with sub-lethal (LD15) and median lethal (LD50) concentrations of emamectin benzoate via topical application. (B) Fecundity of BPH when 4th instar nymphs were treated with sub-lethal (LC15) and median lethal (LC50) concentrations of abamectin via systemic exposure. All data are presented as the mean ± s.e.m. Different lower-case letters above the bars indicate significant differences (One-way ANOVA with Tukey’s Multiple Range Test, p < 0.05).

Earlier version: To examine if EB also stimulates egg-laying in other insect species we conducted bioassays on the small brown planthopper, Laodelphax striatellus, the white backed planthopper, Sogatella furcifera and fruit flies Drosophila melanogaster. In contrast to our findings on BPH, we found that sub-lethal doses (LC15 and LC50) of EB inhibits fecundity of female L. striatellus, (Figure1-figure supplement 2A-C) and has no impact on the fecundity of S. furcifera, (Figure1-figure supplement 2D-F). In addition, we found that sublethal doses (LC15 or LC50) of EB also inhibit fecundity in D. melanogaster (Figure1-figure supplement 2G and H). These results indicate that the stimulation of reproduction by EB in BPH is species-specific and does not extend to even related insect species.

Revised version: To examine if EB also stimulates egg-laying in other insect species we conducted bioassays on the small brown planthopper, Laodelphax striatellus, the white backed planthopper, Sogatella furcifera and fruit flies Drosophila melanogaster. In contrast to our findings in BPHs, we found that sub-lethal doses (LC15 and LC50) of EB inhibits fecundity of female L. striatellus, (Figure1-figure supplement 2A-C) and has no impact on the fecundity of S. furcifera, (Figure1-figure supplement 2D-F). In addition, we found that sublethal doses (LC15, LC30 or LC50) of EB also inhibit fecundity in D. melanogaster (Figure1-figure supplement 2G and H). These results indicate that the stimulation of reproduction by EB in BPH is species-specific and does not even extend to related insect species.

Added in the revised version: Fecundity of small brown planthopper, Laodelphax striatellus, (A-C) white backed planthopper, Sogatella furcifera (D-F) and fruit fly, Drosophila melanogaster (G and H) when larvae and newly emerged adults were treated with sub-lethal concentrations of emamectin benzoate. All data are presented as the mean ± s.e.m. Different lower-case letters above the bars indicate significant differences (One-way ANOVA with Tukey’s Multiple Range Test, p < 0.05).

The impact of EB treatment on BPH reproductive fitness

Earlier version: To better understand the effects of EB on the reproductive fitness of BPH, the preoviposition period, emergence rate, female ratio, female longevity and female weight were evaluated following exposure using the systemic bioassay. The preoviposition period of females treated with the LC50 of EB decreased significantly compared with the control (Figure 1-figure supplement 3A). In contrast no significant effects of EB on emergence rate and female ratio were observed (Figure 1-figure supplement 3B and C). Female survival analysis showed that exposure of 4th instar nymphs to the LC50 of EB has no impact on female longevity (Figure 1-figure supplement 3D). Interestingly, brachypterism (long-wing) female ratio and female weight were significantly increased after EB exposure (Figure 1-figure supplement 3E and F).

Revised version: To better understand the effects of EB on the reproductive fitness of BPH, the duration of the preoviposition period, emergence rate, female ratio, female longevity and female weight were evaluated following exposure using the systemic bioassay. The preoviposition period of females treated with the LC50 of EB decreased significantly compared with the control (Figure 1-figure supplement 3A). In contrast no significant effects of EB on emergence rate and female ratio were observed (Figure 1-figure supplement 3B and C). Female survival analysis showed that exposure of 4th instar nymphs to the LC50 of EB has no impact on female longevity (Figure 1-figure supplement 3D). Interestingly, occurrence of brachypterism (short-wing phenotype) in females and female weight were significantly increased after EB exposure (Figure 1-figure supplement 3E and F). It was reported that short-winged females usually display higher fecundity than long-winged females in BPHs and two different insulin receptors determine the alternative wing morphs (Xu et al., 2015). However, whether insulin signaling is involved in the fecundity increase of BPHs in our experiments needs further investigation.

Added in the revised version: The impact of emamectin benzoate on the reproductive fitness of BPH. Fourth instar nymphs were treated with the LC50 concentration of emamectin benzoate in systemic bioassays. (A) Preoviposition period: Preoviposition refers to the period in an insect’s life cycle between the time it becomes an adult and the time it starts laying eggs. (B) Emergence rate: the rate of emergence of adults; (C) Female ratio: the ratio of female (to male) insects in the total of emerged adults; (D) Female adult longevity. (E) Brachypterism female ratio: the ratio of short-winged to long-winged adults; (F) The weight per female. All data are presented as the mean ± s.e.m. Different lower-case letters above the bars indicate significant differences (Student’s t test, p < 0.05).

Earlier version: EB promotes ovarian development in BPH

Revised version: EB promotes ovarian maturation in BPH

Earlier version: To investigate the cause of increased female weight following EB exposure we examined if EB influences ovarian development in BPH. We dissected and compared the ovaries of females treated with the LC50 of EB at 1, 3, 5 and 7 days after eclosion (DAE) with control females. At 3, 5 and 7 DAE, the number of detained eggs of BPH in the EB treated group were significantly higher than that of controls (Figure 2A and B). We also explored whether EB treatment could enhance or impair oogenesis in BPH. However, dissection of various developmental stages revealed that emamectin benzoate treatment has no significant effects on oogenesis (Figure 2C and D).

Revised version: To investigate the cause of increased egg-laying following EB exposure we examined if EB influences ovarian maturation in BPH. We dissected and compared the ovaries of females treated with the LC50 of EB at 1, 3, 5 and 7 days after eclosion (DAE) with control females. At 3, 5 and 7 DAE, the number of eggs in the ovary of BPH in the EB treated group were significantly higher than that of controls (Figure 2A and B). We also observed that EB treatment of female adults also increases the number of mature eggs in the ovary (Figure 2-figure supplement 1). Hence, we found that EB does not affect normal egg developmental stages. However, our results suggest that EB treatment can promote oogenesis and as a result the insects both produce more eggs in the ovary and more eggs were laid.

Earlier version: The impact of emamectin benzoate on ovarian development in BPH. Fourth instar nymphs were treated with the LC50 concentration of emamectin benzoate in systemic bioassays. (A) Ovarian development in EB treated BPH at 1, 3, 5 and 7 days after eclosion (DAE) compared to untreated controls. Scale bar = 1,000 μm. (B) Number of detained eggs in the ovaries of EB treated BPH females measured at 1, 3, 5 and 7 DAE compared to controls. All data are presented as the mean ± s.e.m. Asterisks indicate values significantly different from the control using student t test (ns, no significant; *p < 0.05 and **p < 0.01). (C) Different developmental stages of BPH eggs. (D) No impairment of emamectin benzoate on oogenesis of BPH. Scale bar = 100 μm.

Revised version: The impact of emamectin benzoate on ovarian maturation in BPH. Fourth instar nymphs were treated with the LC50 concentration of emamectin benzoate in systemic bioassays. (A) Ovarian development in EB treated BPH at 1, 3, 5 and 7 days after eclosion (DAE) compared to untreated controls. Scale bar = 1,000 μm. (B) Number of mature eggs in the ovaries of EB treated 4th instar BPH nymphs measured at 1, 3, 5 and 7 DAE compared to controls. All data are presented as the mean ± s.e.m. Asterisks indicate values significantly different from the control using student t test (ns, no significant; *p < 0.05 and **p < 0.01). (C) Different developmental stages of BPH eggs. (D) No impairment of emamectin benzoate on oogenesis of BPH. Scale bar = 100 μm.

Added in the revised version: Number of mature eggs in the ovaries of EB treated BPH female adults compared to controls. The data are presented as the mean ± s.e.m. Asterisks indicate values significantly different from the control using student t test (*p < 0.05).

Added in the revised version: Unlike the model insect Drosophila melanogaster, studies on oogenesis in the BPH are relatively limited. Therefore, we used immunohistochemistry to stain eggs at different stages of BPH development. We captured images with a laser confocal microscope, to document each stage of oogenesis in this species. We found that oogenesis in the BPH differs from that in D. melanogaster and belongs to the telotrophic meroistic type, where the oocyte is connected to nutrient-providing cells in the anterior trophic region of the ovariole through a nutritive cord. Oogenesis in the BPHs occurs in three main compartments (Figure 2-figure supplement 2):

Added in the revised version: Terminal Filament (TF): This region, composed of terminal filament cells, is similar to the terminal filament structure found in Drosophila.

Added in the revised version: Germarium: This region is where germline stem cells (GSCs) differentiate into oocytes and follicle cells. It typically consists of cap cells adjacent to the terminal filament, germline stem cells, escort cells, and cystoblasts.

Added in the revised version: Egg Chamber: The egg chamber in the BPH primarily contains eggs at six developmental stages. The distinct layering of follicle cells, visible through cytoskeletal staining, is the main feature distinguishing the six different developmental stages listed in the figure legend (see also Figure 2-figure supplement 2).

Added in the revised version: We next explored whether EB treatment could enhance or impair oogenesis in BPH. However, dissection of various developmental stages revealed that emamectin benzoate treatment has no significant negative effects on developmental stages and morphology of oogenesis in BPH (Figure 2C and D). Considering the number of eggs laid by EB treated females was larger than in control females (Figure 1 and Figure 1-figure supplement 1), our data indicated that EB treatment of BPHs can both promote oogenesis and fecundity (egg-laying) in BPHs.

Added in the revised version: The oogenesis of BPHs. See the main text for detailed information.

Added in the revised version: Stage 1: Connected to the germarium, the egg is trapezoidal. Stage 2: The egg is larger than in Stage 1 and is square to rectangular. Stage 3: The egg begins to elongate toward both poles, taking on an oval shape. Stage 4: The egg continues to enlarge and elongate, becoming spindle-shaped. Due to rapid growth, follicle cells on the egg surface start dividing, indicating nuclear division. Stage 5: The egg further enlarges, and most follicle cells display binucleation, the most notable feature of this stage. Small amounts of lipids start entering the egg. Stage 5 is the final stage before maturation, after which the egg enters the pedicel, a structure primarily formed by muscle fibers. The expanded mature egg stretches the muscle fibers in the pedicel, increasing the spacing between them. Stage 6: This is the mature egg stage, which occurs after the egg passes through the pedicel. During this stage, the egg undergoes rapid growth, continuously absorbing vitellogenin and lipids (Nile red staining confirmed that Stage 6 has the highest lipid content). This absorption causes the egg’s surface color to transition from the transparency seen in Stage 5 to mature eggs’ dark, opaque appearance. Once the egg enters the pedicel and becomes completely opaque, it reaches full maturity. The mature egg then moves into the lateral oviduct for oviposition, completing the egg-laying process.

Earlier version: EB exposure enhances the abundances of storage macromolecules and circulating sugars in BPH

Revised version: EB exposure increases circulating sugars and the storage of macromolecules in BPHs

Earlier version: Nutritional status is an important indicator of reproductive fitness. Thus, to investigate whether EB affects intermediary metabolism and energy storage of BPH, glycogen, triacylglyceride (TAG), total protein content, cholesterol and four circulating carbohydrates were quantified in 4th instar BPH nymphs following exposure to the LC50 of EB.

Revised version: Nutritional status is an important indicator of reproductive fitness. Thus, to investigate whether EB affects intermediary metabolism and energy storage of BPHs, glycogen, triacylglyceride (TAG), total protein content, cholesterol and four circulating carbohydrates were quantified in 4th instar BPH nymphs following exposure to the LC50 of EB.

Earlier version: We found that EB exposure has no impact on glycogen levels (Figure 2-figure supplement 1A). The amounts of TAG in EB-treated BPH were 27% higher (p < 0.05) than those in controls, but only in BPH of the late fifth instar (5L) stage, with no significant differences observed in subsequent developmental stages (Figure 2-figure supplement 1B). The amount of total protein content in EB-treated BPH was higher than the control groups in the case of all developmental stages from 5L nymph to 7DAE (Figure 2-figure supplement 1C). EB exposure also increased cholesterol levels at 4 and 5 DAE (Figure 2-figure supplement 1D). Compared with the solvent control, EB treatment caused significant increases (p < 0.05) in the levels of sucrose, glucose, fructose, and trehalose (Figure 2-figure supplement 1E-H). Thus, collectively, these data provide evidence that EB exposure leads to energy mobilization and the metabolism of carbohydrates and lipids in BPH.

Revised version: We found that EB exposure has no impact on glycogen levels (Figure 2-figure supplement 2A). The amount of TAG in EB-treated BPHs was 27% higher (p < 0.05) than those in controls, but only in BPH of the late fifth instar (5L) stage, with no significant differences observed in subsequent developmental stages (Figure 2-figure supplement 2B). The amount of total protein content in EB-treated BPH was higher than the control groups in the case of all developmental stages from 5L nymph to 7DAE (Figure 2-figure supplement 2C). EB exposure also increased cholesterol levels at 4 and 5 DAE (Figure 2-figure supplement 2D). Compared with the solvent control, EB treatment caused significant increases (p < 0.05) in the levels of sucrose, glucose, fructose, and trehalose (Figure 2-figure supplement 2E-H). Thus, collectively, these data provide evidence that EB exposure leads to energy mobilization and the metabolism of carbohydrates and lipids in BPHs.

Added in the revised version: Amounts of Glycogen (A), TAG (B), total protein content (C), cholesterol (D) and four circulating sugars including sucrose, glucose, fructose and trehalose (E-H) after BPH exposure to EB. All data are presented as the mean ± s.e.m. The differences between the EB-treated and solvent-treated BPH were analyzed using unpaired student t-test (*, p < 0.05; **, p < 0.01; ***, p < 0.001; ****, p < 0.0001).

EB stimulates egg-laying that is mediated by the JH signaling pathway

Earlier version: Given the important role of juvenile hormone (JH) in vitellogenesis and egg development in insects (19, 25, 30–32), we asked whether EB-treatment could influence the titer of JH in BPH. As measured by ELISA, the juvenile hormone titer of BPH nymphs treated with the LC50 concentration of EB was significantly lower than that of controls in systemic bioassays during the middle and late stages of the 4th instar (Figure 3A). However, at 2, 3 and 4 DAE, the JH titer in the EB treated group was significantly higher than that of the control (Figure 3A). Interestingly, the titer of another important insect hormone, the steroid ecdysone, was not significantly different between EB-treated BPH and solvent-treated controls (Figure 3-figure supplement 1). To independently validate the results of ELISA assays, we employed HPLC-MS/MS to measure JH titer in BPH following EB exposure (30, 33, 34). The results showed that JH III titer significantly decreased after EB-treatment at the late 4th instar nymph stage (Figure 3A and B), but significantly increased at the third day after eclosion (3 DAE) (Figure 3A and C). To further investigate the role of JH in EB- enhanced fecundity in BPH, we treated BPH with methoprene and pyriproxyfen, JH analogues or biologically active JH mimics respectively, to determine whether they can stimulate fecundity in BPH. Both compounds significantly increased egg-laying in BPH (Figure 3D). Takes together these results reveal that EB stimulates an increase in JH titer that induces fecundity in BPH. Since we found that EB could induce JH synthesis in the BPH, we asked whether EB could influence the expression of genes that are involved in JH synthesis or degradation. For this we treated 4th instar nymphs of BPH with the LC50 concentration of EB using systemic bioassays, and then collected early (5E), middle (5M) and late (5L) stage of 5th instar nymph and 1-7 DAE female adults (DAE) for analysis. Quantitative PCR was then used to examine the expression of key genes previously implicated in the regulation of JH (35, 36).

Revised version: Given the important role of juvenile hormone (JH) in vitellogenesis and egg development in insects (Jing et al., 2021; Ling and Raikhel, 2021; Luo et al., 2021; Santos et al., 2019; Zheng et al., 2022), we asked whether EB-treatment could influence the titer of JH in BPH. As measured by ELISA, the juvenile hormone titer of BPH nymphs treated with the LC50 concentration of EB was significantly lower than that of controls in systemic bioassays during the middle and late stages of the 4th instar (Figure 3A). However, at 2, 3 and 4 DAE, the JH titer in the EB treated group was significantly higher than that of the control (Figure 3A). Interestingly, the titer of another important insect hormone, the steroid ecdysone, was not significantly different between EB-treated BPH and solvent-treated controls (Figure 3-figure supplement 1). To independently validate the results of ELISA assays, we employed HPLC-MS/MS to measure JH titer in BPH following EB exposure (Bownes and Rembold, 1987; Guo et al., 2020; Jing et al., 2021). The results showed that the JH III titer significantly decreased after EB-treatment at the late 4th instar nymph stage (Figure 3A and B), but significantly increased at the third day after eclosion (3 DAE) (Figure 3A and C). We also investigated the effects of EB treatment on the JH titer of female adults. The data indicates that the JH titer was also significantly increased in the EB-treated female adults compared with controls (Figure 3-figure supplement 3A). However, the steroid ecdysone, was also not significantly different between EB-treated BPH and controls (Figure 3-figure supplement 3B).

Earlier version: EB induced reproduction in BPH is mediated by components of the JH signaling pathway.

Revised version: EB induced reproduction in BPHs is mediated by components of the JH signaling pathway.

(A) The titer of JH III (as measured by ELISA assay) at different developmental stages of BPH when 4th instar nymphs were treated with the median lethal concentrations of EB. (B and C) The titer of JH III (as measured by HPLC-MS/MS) in BPH females at 4L and 3 DAE when treated with median lethal concentrations of EB. (D) Oviposition rate of BPH when 4th instar nymphs were treated with 4 ppm methoprene or 10 ppm pyriproxyfen. (E-J) Expression of selected JH-related genes (FAMeT, JHAMT, Met, Kr-h1, Vg, and JHE) in EB-treated BPH. (K) Egg production following silencing of JHAMT with or without EB application. (L) Egg production following silencing of met with or without EB application. (M) Egg production after silencing Kr-h1 with or without EB application. (N) Schematic illustrating the proposed impact of EB on the JH signaling pathway leading to enhanced reproduction. The question mark indicates one or more potential additional signals. All data are presented as means ± s.e.m. Student’s t test was used to compare the two samples. One-way ANOVA with Tukey’s multiple comparisons test was used to compare more than two samples. ns, no significant difference; Asterisks indicate values significantly different from the control (ns, no significant; *p < 0.05, **p < 0.01, ***p < 0.001 and ****p < 0.0001). Different lower-case letters above the bars indicate significant differences (p < 0.05).

Added in the revised version: Ecdysone titer at different developmental stages of BPH when 4th instar nymphs were treated with median lethal concentrations of EB.

Added in the revised version: (A) Expression of Kr-h1 following RNAi knockdown. (B) Expression of Vg when Kr-h1 was silenced in BPH. All data are presented as means ± s.e.m. *p < 0.05; Mann–Whitney test.

Added in the revised version: (A) The titer of JH III (as measured by HPLC-MS/MS) in BPH females at adult stages when treated with median lethal concentrations of EB. (B) The titer of ecdysone titer of BPH when female adults were treated with median lethal concentrations of EB. (C and D) Expression of JHAMT and Kr-h1 in EB-treated BPH female adults. All data are presented as means ± s.e.m. Student’s t test was used to compare the two samples. Asterisks indicate values significantly different from the control (ns, no significant; *p < 0.05).

Added in the revised version: To further investigate the role of JH in EB-enhanced fecundity in BPHs, we treated BPHs with methoprene and pyriproxyfen, JH analogues or biologically active JH mimic respectively, to determine whether they can stimulate fecundity in BPH. Both compounds significantly increased egg-laying in BPHs (Figure 3D). Taken together these results reveal that EB stimulates an increase in JH titer and that JH signaling enhances fecundity in BPHs.

Added in the revised version: Since we found that EB could induce JH synthesis in the BPH, we asked whether EB could influence the expression of genes that are involved in JH biosynthesis or degradation. For this we treated 4th instar nymphs or newly emerged female adults of BPHs with the LC50 concentration of EB using systemic bioassay, and then collected early (5E), middle (5M) and late (5L) stage of 5th instar nymph and 1-7 DAE female adults for analysis. Quantitative PCR was then used to examine the expression of key genes previously implicated in the regulation of JH (Gospocic et al., 2021; Nouzova et al., 2018).

Added in the revised version: Farnesoic acid O-methyltransferase (FAMeT) (Liu et al., 2010), is known as an important enzyme in the JH biosynthetic pathway, catalyzing methylation of farnesoic acid (FA) to methyl farnesoate (MF) (Nouzova et al., 2018). We found that this gene was significantly upregulated in 5M instar nymph to 4 DAE after EB-treatment (1.0-fold to 3.0-fold) (Figure 3E). Juvenile hormone acid methyltransferase (JHAMT) (Nouzova et al., 2018; Shinoda and Itoyama, 2003), which is involved in the biosynthesis of JH, was also upregulated in EB-treated BPH at 1 to 5 DAE compared to controls (1.5-fold to 3.0-fold) (Figure 3F). We furthermore found that EB treatment in female adults stimulates JHAMT expression (Figure 3-figure supplement 3C).

Added in the revised version: Methoprene-tolerant (Met), belongs to the basic helix–loop–helix Per/Arnt/Sim (bHLH-PAS) family of transcription factors and is the intracellular (nuclear) receptor of JH (Riddiford, 2020; Zhu et al., 2019). The levels of met mRNA increased in EB-treated BPH at the 5M and 5L instar nymph and 1 to 5 DAE adult stages compared to controls (1.7-fold to 2.9-fold) (Figure 3G).

Added in the revised version: Krüppel homolog 1 (Kr-h1), a transcriptional target of JH signaling, is reported to be sensitive to levels of JH and its expression levels are directly correlated with JH titers (Gospocic et al., 2021; Jindra et al., 2013; Zhang et al., 2022a). We found that Kr-h1 was significantly upregulated in the adults of EB-treated BPH at the 5M, 5L nymph and 4 to 5 DAE stages (4.7-fold to 27.2-fold) when 4th instar nymph and female adults were treated with EB (Figure 3H and Figure 3-figure supplement 3D).

Added in the revised version: Similarly, the expression levels of vitellogenin (Vg), a key downstream component of JH signaling triggering ovarian development in insects including BPH (Shen et al., 2019), was markedly increased in females at 2–5 DAE by EB (1.7-fold to 5.5-fold) (Figure 3I). During the maturation process, eggs require uptake of vg, and an increase in vg content can accelerate egg maturation, producing more mature eggs. Our molecular data suggest that EB treatment leads to an upregulation of vg expression. Based on these findings, we conclude that the increase in egg-laying caused by EB treatment is due to the upregulation of vg (Figure 3I), which raises Vg content, promoting the uptake of Vg by maturing eggs and resulting in the production of more mature eggs.

Added in the revised version: Juvenile hormone esterase (JHE) is the primary JH-specific degradation enzyme that plays a key role in regulating JH titers (Kamita et al., 2003). Interestingly, we observed a significant upregulation of JHE mRNA levels in the early and middle 5th instar nymph stage followed by downregulation in 5L instar nymphs of EB-treated BPH (1.3-fold to 2.6-fold) (Figure 3J).

Added in the revised version: In combination these results reveal that EB has a profound impact on the expression of key genes involved in the biosynthesis of JH or in downstream signaling pathway genes that might promote egg development and female fecundity.

Removed in the revised version: Farnesoic acid O-methyltransferase (FAMeT) (37), is known as an important enzyme in the JH biosynthetic pathway, catalyzing methylation of farnesoic acid (FA) to methyl farnesoate (MF) (36). We found that this gene was significantly upregulated in 5M instar nymph to 4 DAE after EB- treatment (1.0-fold to 3.0-fold) (Figure 3E). Another gene, juvenile hormone acid methyltransferase (JHAMT) (36, 38), which is involved in the biosynthesis of JH, was also upregulated in EB-treated BPH at 1 to 5 DAE than controls (1.5-fold to 3.0-fold) (Figure 3F). Methoprene-tolerant (Met), belongs the basic helix–loop–helix Per/Arnt/Sim (bHLH-PAS) family of transcription factors and is the intracellular (nuclear) receptor of JH (39, 40). The levels of met mRNA increased in EB-treated BPH at the 5M and 5L instar nymph and 1 to 5 DAE stages compared to controls (1.7-fold to 2.9-fold) (Figure 3G). Krüppel homolog 1 (Kr-h1), a transcriptional target of JH signaling, is reported to be sensitive to levels of JH and its expression levels are directly correlated with JH titers (35, 41, 42). We found that Kr-h1 was significantly upregulated in the adults of EB-treated BPH at the 5M, 5L nymph and 4 to 5 DAE stages (4.7-fold to 27.2-fold) (Figure 3H). Similarly, the expression levels of vitellogenin (Vg), a key downstream component of JH signaling triggering ovarian development in insects including BPH (43), was markedly increased in females at 2–5 DAE by EB (1.7-fold to 5.5-fold) (Figure 3I). Juvenile hormone esterase (JHE) is the primary JH-specific degradation enzyme that plays a key role in regulating JH titers (44). Interestingly, we observed a significant upregulation of JHE mRNA levels in the early and middle 5th instar nymph stage followed by downregulation in 5L instar nymphs of EB-treated BPH (1.3-fold to 2.6-fold) (Figure 3J). In combination these results reveal that EB has profound impacts on the expression of key genes involved in the synthesis of JH or downstream signaling pathway genes that might promote egg development.

Earlier version: To further understand whether these JH pathway-related genes were involved in egg-laying behavior in BPH, we performed RNAi experiments to downregulate the expression of JHAMT, Met and Kr-h1 (Figure 3-figure supplement 2A-C). We found that silencing of these three genes downregulated the titer of JH in BPH providing functional evidence of their role in the regulation of JH (Figure 3-figure supplement 2D). Furthermore, silencing of Kr-h1 inhibited FAMeT and Vg gene expression while increasing JHE gene expression (Figure 3-figure supplement 2E-H). Importantly, silencing of JHAMT, Met and Kr-h1 gene in female BPH was also found to suppress egg-laying (Figure 3K-M). However, this phenotype was rescued by EB treatment, which restored egg-laying to normal levels in BPH injected with JHAMT, Met and Kr-h1 dsRNA (Figure 3K-M). Together these results provide a mechanistic understanding of how EB enhances fecundity in BPH by modulating the expression of key genes involved in JH synthesis (Figure 3N).

Revised version: To further understand whether these JH pathway-related genes were involved in egg-laying behavior in BPH, we performed RNAi experiments to downregulate the expression of Kr-h1 (Figure 3-figure supplement 2A). We found that silencing of Kr-h1 diminished vg gene expression (Figure 3-figure supplement 2B). Importantly, diminishment of JHAMT, Met and Kr-h1 genes in female BPHs was also found to suppress egg-laying (Figure 3K-M). However, this phenotype was rescued by EB treatment, which restored egg-laying to normal levels in BPH injected with JHAMT, Met and Kr-h1 dsRNA (Figure 3K-M). Note that this rescue is possible since knockdown of the genes is incomplete when using dsRNA injection (and residual gene expression allows for EB action). Together these results provide a mechanistic understanding of how EB enhances fecundity in BPH by modulating the expression of key genes involved in JH synthesis (Figure 3N).

EB induces JH biosynthesis through the peptidergic AstA/AstAR signaling pathway

Earlier version: The timing and level of JH biosynthesis can be precisely regulated by the neuropeptides, stimulatory allatotropins (ATs) and inhibitory allatostatins (Asts), in many insects (45–52). Insects can, in a species-specific manner, produce one type of AT and three types of Asts: FGL-amide Ast (AstA) (46, 53); myoinhibitory peptide (MIP or AstB) and PISCF Ast (AstC) (50, 54). In some species, there also exists two paralogue genes of AstCs which are named AstCC and AstCCC (55, 56). Interestingly, the allatostatic activity of these three types of Ast peptides varies between insect species so that in each species only one type of Ast (for example AstA) controls JH production (50, 51, 54, 57).

Revised version: The timing and level of JH biosynthesis can be precisely regulated by neuropeptides, stimulatory allatotropins (ATs) and inhibitory allatostatins (Asts), in many insects (Bellés et al., 1999; Kataoka et al., 1989; Kramer et al., 1991; Lorenz et al., 1995; Stay and Tobe, 2007; Verlinden et al., 2015; Wegener and Chen, 2022; Woodhead et al., 1989). Insects can, in a species-specific manner, produce one type of AT and three types of Asts: FGL-amide Ast (AstA) (Woodhead et al., 1989; Yin et al., 2006); myoinhibitory peptide (MIP or AstB) (Stay and Tobe, 2007) and PISCF Ast (AstC) (Wang et al., 2012). In some species, there also exists two paralogue genes of AstCs which are named AstCC and AstCCC (Veenstra, 2009, 2016). Interestingly, the allatostatic activity of these three types of Ast peptides varies between insect species so that in each species only one type of Ast (for example AstA) controls JH production (Nässel, 2002; Stay and Tobe, 2007; Wang et al., 2012; Wegener and Chen, 2022).

Earlier version: Analysis of the BPH neural transcriptome sequence data has revealed the presence of one AT, four types of Asts and four corresponding receptors, allatotropin receptor (A16, ATR), AstA receptor (A2, AstAR), AstB (MIP) receptor (A10, AstBR or MIPR) and AstC receptor (A1, AstCR) (58). We cloned the five neuropeptide genes (AT, AstA, AstB/MIP, AstCC and AstCCC) and confirmed the sequence obtained from transcriptome data (Figure 4-figure supplement 1) (58). Interestingly, we found that AstC is missing in the genome of BPH and only AstCC and AstCCC are present (Figure 4-figure supplement 1). Next, we also cloned their corresponding receptors (55) including ATR (A16), AstAR (A2), AstBR (A10) and AstCR (A1) which might be activated by AstCC and/or AstCCC (42, 55, 59). Sequence alignments and phylogenetic analysis are shown in Figure 4-figure supplement 2.

Revised version: Analysis of the BPH neural transcriptome sequence data has revealed the presence of one AT, four types of Asts and four corresponding receptors, allatotropin receptor (A16, ATR), AstA receptor (A2, AstAR), AstB (MIP) receptor (A10, AstBR or MIPR) and AstC receptor (A1, AstCR) (Tanaka et al., 2014). We cloned the five neuropeptide genes (AT, AstA, AstB/MIP, AstCC and AstCCC) and confirmed the sequence obtained from transcriptome data (Figure 4-figure supplement 1) (Tanaka et al., 2014). Interestingly, we found that AstC is missing in the genome of BPH and only AstCC and AstCCC are present (Figure 4-figure supplement 1). Next, we also cloned their corresponding receptors (Veenstra, 2009) including ATR (A16), AstAR (A2), AstBR (A10) and AstCR (A1) which might be activated by AstCC and/or AstCCC (Audsley et al., 2013; Veenstra, 2009; Zhang et al., 2022a). Sequence alignments and phylogenetic analysis are shown in Figure 4-figure supplement 2.

Removed in the revised version: EB induced reproduction in BPH is mediated by the AstA/AstAR and JH signaling pathway.

Removed in the revised version: (A and B) Expression of AstA and AstAR in different stages of BPH following EB treatment. (C) Downregulation of AstAR using RNAi leads to a reduction in mRNA expression level. (D) Egg production in female BPH following silencing of AstAR gene. (E-I) Expression of selected JH signaling pathway related genes (FAMeT, JHAMT, Met, Kr-h1 and JHE) in AstAR silenced BPH. (J) JHIII titer of BPH females after AstAR gene silencing determined by HPLC-MS/MS. (K) Number of eggs laid per female in 48h following injection of six AstA1-AstA6 mature peptides and one AT mature peptide. Fifty nanoliter of PBS (as control) and seven peptides (20 pmol/insect) were injected into female BPH three days after eclosion. (L and M) The JH III titer of BPH females at different time points following AstA or AT injection. (N) Schematic of the proposed role of AstA/AstAR in the regulation of JH following EB exposure. The question mark indicates one or more possible additional signals. All data are presented as means ± s.e.m. Student’s t test was used to compare the two samples. One-way ANOVA with Tukey’s multiple comparisons test was used to compare more than two samples. ns, no significant difference; Asterisks indicate values significantly different from the control (ns, no significant; *p < 0.05, **p < 0.01, ***p < 0.001 and ****p < 0.0001). Different lower-case letters above the bars indicate significant differences (p < 0.05).

Earlier version: Quantitative PCR was then used to examine if EB treatment could influence the expression of the genes encoding these neuropeptides and their receptors. Treating BPH with the LC50 concentration of EB significantly increased the expression of AT, ATR and AstCCC while resulting in the downregulation of AstA, AstB/mip, AstCC, AstAR and AstBR/mipr at the adult stage (Figure 4A and B, Figure 4-figure supplement 3A-G). Among these, AstA and AstAR were the most downregulated genes after EB treatment (Figure 4A and B, Figure 4-figure supplement 3H) and thus the AstA/AstAR signaling system was selected for subsequent functional analysis. Silencing of AstAR in female BPH using RNAi (Figure 4C), resulted in an increased number of eggs laid per female compared with dsgfp-injected controls (Figure 4D). Interestingly, silencing AstAR also resulted in the upregulation of FAMeT, JHAMT, Met and Kr-h1 which are involved in the JH biosynthesis/signaling (Figure 4E-H). However, JHE was not influenced by AstAR silencing (Figure 4I). We therefore investigated whether silencing the AstAR gene could influence JH titer in BPH. A significantly increased JH titer was observed in AstAR silenced BPH compared with controls (Figure 4J). Thus, our data strongly suggest that AstA is a key inhibitor of JH production in BPH.

Revised version: Quantitative PCR was then used to examine if EB treatment could influence the expression of the genes encoding these neuropeptides and their receptors. Treating 4th instar or female adults BPH with the LC50 concentration of EB significantly increased the expression of AT, ATR and AstCCC while resulting in the downregulation of AstA, AstB/mip, AstCC, AstAR and AstBR/mipr at the adult stage (Figure 4A and B, Figure 4-figure supplement 3A-G). Among these, AstA and AstAR were the most downregulated genes after EB treatment (Figure 4A and B, Figure 4-figure supplement 3H) and thus the AstA/AstAR signaling system was selected for subsequent functional analysis. Silencing of AstAR in female BPH using RNAi (Figure 4C), resulted in an increased number of eggs laid per female compared with dsgfp-injected controls (Figure 4D). Interestingly, silencing AstAR also resulted in the upregulation of JHAMT, Met and Kr-h1 which are involved in the JH biosynthesis/signaling (Figure 4E-H). However, JHE was not influenced by AstAR silencing (Figure 4I). We next investigated whether silencing the AstAR gene would influence the JH titer in BPHs. A significantly increased JH titer was observed in AstAR silenced BPH compared with controls (Figure 4J). Thus, our data strongly suggest that AstA is a key inhibitor of JH production in BPH.

Added in the revised version: EB induced reproduction in BPHs is mediated by the AstA/AstAR and JH signaling pathway.

Added in the revised version: (A and B) Expression of AstA and AstAR in different stages of BPH following EB treatment. (C) Downregulation of AstAR using RNAi leads to a reduction in mRNA expression level. (D) Egg production in female BPH following silencing of AstAR gene. (E-I) Expression of selected JH signaling pathway related genes (JHAMT, Met, Kr-h1 and JHE) in AstAR silenced BPH. (J) JHIII titer of BPH females after AstAR gene silencing determined by HPLC-MS/MS. (K) Number of eggs laid per female in 48h following injection of the six mature AstA1-AstA6 peptides and one mature AT peptide. Fifty nanoliter of PBS (as control) and seven different peptides (20 pmol/insect) were injected into female BPHs three days after eclosion. (L and M) The JH III titer of BPH females at different time points following AstA or AT injection. (N) Schematic of the proposed role of AstA/AstAR in the regulation of JH following EB exposure. The question mark indicates one or more possible additional signals. All data are presented as means ± s.e.m. Student’s t test was used to compare the two samples. One-way ANOVA with Tukey’s multiple comparisons test was used to compare more than two samples. ns, no significant difference; Asterisks indicate values significantly different from the control (ns, no significant; *p < 0.05, **p < 0.01, ***p < 0.001 and ****p < 0.0001). Different lower-case letters above the bars indicate significant differences (p < 0.05).

Added in the revised version: Alignments of the amino acid sequences of: (A) AT, (B) AstA, (C) AstB/MIP, (D) AstCC and (E) AstCCC peptides from select species. AT, AstA, AstB/MIP and AstCCC are predicted to have a C-terminal amide. The mature peptides belonging to the same species have been highlighted with the same color. Species names are as follows: Nillu (Nilaparvata lugens), Locmi (Locusta migratoria), Scham (Schistocerca americana), Homvi (Homalodisca vitripennis), Manse (Manduca sexta), Spofr (Spodoptera frugiperda), Drome (Drosophila melanogaster), Spoex (Spodoptera exigua), Nasvi (Nasonia vitripennis), Grybi (Gryllus bimaculatus), Bommi (Bombyx mori); Mesma (Mesobuthus martensii), Stear (Stegodyphus araneomorph), Limpo (Limulus polyphemus), Carma (Carcinus maenas), Strma (Strigamia maritima), Athro (Athalia rosae), Apime (Apis mellifera), Dapma (Daphnia magna). Black lines under the sequences indicate the locations of the disulfide bridges in the mature peptides. The accession numbers of the sequences are listed in Figure 4-figure supplement 1 source data.

Added in the revised version: Phylogenetic tree of the predicted BPH (*) allatotropin receptor (A16, ATR), allatostatins A receptor (A2, AstAR), AstB (MIP) receptor (A10, AstBR or MIPR) and allatostatins C receptor (A1, AstCR) with other insect species. The tree was generated using the maximum likelihood method. Drosophila melanogaster metabotropic glutamate receptor was included as an outgroup. The accession numbers of the sequences used for this phylogenetic tree are listed in Figure 4-figure supplement 2 source data.

Added in the revised version: EB induced changes in the expression of AT, AstB, AstCC, AstCCC, ATR, AstBR and AstCR in BPH.

Added in the revised version: All data are presented as means ± s.e.m. Student’s t test was used to compare the two samples. ns, no significant difference; Asterisks indicate values significantly different from the control (*p < 0.05, **p < 0.01, and ***p < 0.001).

Earlier version: Finally, we investigated whether injection of mature Ast and AT peptides could influence the number of eggs laid and JH titer in BPH. We synthesized one AT, six AstAs (AstA1 to AstA6), one AstCC and one AstCCC peptide according to their determined sequences (Figure 4-figure supplement 1). Indeed, we found that injection of AstA1 and AstA6 reduced the number of eggs laid per female in 48 h (Figure 4K). We also observed that AstA1 injection after 16h and 48h decreased the JH titer significantly (Figure 4L), and AT injection increased the JH titer after 2h but levels returned to normal 4h after injection (Figure 4M). Collectively, these results provide compelling evidence that EB induces reproduction in BPH through the AstA/AstAR and JH signaling pathways (Figure 4N) and further supports the role of AstA and AT in regulation of JH titer.

Revised version: Finally, we investigated whether injection of mature Ast and AT peptides could influence the number of eggs laid and JH titer in BPH. We synthesized one AT, six AstAs (AstA1 to AstA6), one AstCC and one AstCCC peptide according to their determined sequences (Figure 4-figure supplement 1). Indeed, we found that injection of AstA1 and AstA6 in female adults reduced the number of eggs laid per female over 48 h (Figure 4K). We also observed that AstA1 injection decreased the JH titer as measured 16h and 48h after injection (Figure 4L), and AT injection increased the JH titer after 2h, but levels returned to normal 4h after injection (Figure 4M). Collectively, these results provide compelling evidence that EB induces reproduction in BPH through the AstA/AstAR and JH signaling pathways (Figure 4N) and further supports the role of AstA and AT in regulation of JH titer in this species.

Earlier version: EB-enhanced fecundity in BPH is dependent on its molecular target protein the GluCl

Revised version: EB-enhanced fecundity in BPHs is dependent on its molecular target protein, the GluCl channel

Earlier version: EB and abamectin, are allosteric modulators, which target glutamate-gated chloride channels (GluCls) (60–62). Hence, we examined whether EB-stimulated fecundity in BPH is influenced by its molecular target GluCl. The full length GluCl coding sequence from BPH was cloned and sequenced (Figure 5-figure supplement 1) and the impact of EB on GluCl gene expression examined using quantitative PCR. Treatment of BPH with the LC50 concentration of EB significantly downregulated GluCl gene expression at the 5E and 5M nymph stages while upregulating GluCl gene expression at 2 DAE and 5 DAE in the adult stage (Figure 5A). To examine the role of GluCl gene in BPH fecundity, RNAi was used to knockdown expression of this gene in female BPH (Figure 5B). A significant decrease in the number of eggs laid by per female was observed in dsGluCl-injected insects compared with dsgfp-injected insects (Figure 5C). However, treatment with EB was found to rescue the decreased egg-laying phenotype induced by dsGluCl injection (Figure 5C). To investigate the mechanism by which GluCl expression modulates fecundity we examined if silencing GluCl influences JH titer and JH-related gene expression. Indeed, we observed that RNAi knockdown of GluCl leads to a decrease in JH titer (Figure 5D) and down-regulation of genes including JHAMT which is responsible for JH synthesis, and the JH signaling downstream genes Met and Kr-h1 (Figure 5E-G). In contrast, expression of FAMeT and JHE were not changed in the GluCl silencing insects (Figure 5H and I). We also examined whether silencing GluCl impacts the AstA/AstAR signaling pathway. Silencing GluCl was found to have no impact on the expression of AT, AstA, AstB, AstCC, AstAR, and AstBR. However, the expression of AstCCC and AstCR was significantly upregulated in dsGluCl-injected insects (Figure 5-figure supplement 2A-H). These results suggest that EB activates GluCl which induces JH biosynthesis and release, which in turn stimulates reproduction in BPH (Figure 5J).

Revised version: EB and abamectin, are allosteric modulators, which target glutamate-gated chloride channels (GluCls) (Sparks et al., 2020; Sparks et al., 2021; Wu et al., 2017). Hence, we examined whether EB-stimulated fecundity in BPH is mediated by its molecular target GluCl. The full length GluCl coding sequence from BPH was cloned and sequenced (Figure 5-figure supplement 1) and the impact of EB on GluCl gene expression examined using quantitative PCR. Treatment of BPH with the LC50 concentration of EB significantly downregulated GluCl gene expression at the 5E and 5M nymph stages, while upregulating GluCl gene expression at 2 DAE and 5 DAE in the adult stage (Figure 5A). We also found that EB treatment in female adults also increases the GluCl gene expression (Figure 5-figure supplement 3). To examine the role of GluCl gene in BPH fecundity, RNAi was used to knockdown expression of this gene in female adult BPH (Figure 5B). A significant decrease in the number of eggs laid by per female was observed in dsGluCl-injected insects compared with dsgfp-injected insects (Figure 5C). However, treatment with EB was found to rescue the decreased egg-laying phenotype induced by dsGluCl injection (Figure 5C). Note that this rescue is possible since knockdown of the genes is incomplete when using dsRNA injection (and residual gene expression allows for EB action). To investigate the mechanism by which GluCl expression modulates fecundity we examined if silencing GluCl influences JH titer and JH-related gene expression. Indeed, we observed that RNAi knockdown of GluCl leads to a decrease in JH titer (Figure 5D) and down-regulation of genes including JHAMT which is responsible for JH synthesis, and the JH signaling downstream genes Met and Kr-h1 (Figure 5E-G). In contrast, expression of JHE was not changed in the GluCl knockdown insects (Figure 5H and I). We also examined whether silencing GluCl impacts the AstA/AstAR signaling pathway in female adults. Silencing GluCl in female adults was found to have no impact on the expression of AT, AstA, AstB, AstCC, AstAR, and AstBR. However, the expression of AstCCC and AstCR was significantly upregulated in dsGluCl-injected insects (Figure 5-figure supplement 2A-H). These results suggest that EB activates GluCl which induces JH biosynthesis and release, which in turn stimulates reproduction in BPH (Figure 5J).

EB induced reproduction in brown planthoppers through its molecular target protein GluCl.

Earlier version: (A) Expression of GluCl in EB-treated and untreated BPH. (B) Expression of GluCl following injection of dsGluCl in BPH. (C) Egg production after GluCl gene knockdown in EB-treated and untreated BPH. (D) The JH III titer of BPH females after GluCl gene silencing as quantified using the ELISA method. (E-I) Expression patterns of selected JH-related genes (JHAMT, Met, Kr-h1, FAMeT and JHE) in GluCl silenced BPH. (J) Schematic of the proposed role of GluCl as a molecular target of EB and EB-enhanced reproduction in BPH. The question mark indicates one or more possible additional signals. All data are presented as means ± s.e.m. Student’s t test was used to compare the two samples. One-way ANOVA with Tukey’s multiple comparisons test was used to compare more than two samples. ns, no significant difference; Asterisks indicate values significantly different from the control (ns, no significant; *p < 0.05 and **p < 0.01). Different lower-case letters above the bars indicate significant differences (p < 0.05).

Revised version: (A) Expression of GluCl in EB-treated and untreated BPH. (B) Expression of GluCl following injection of dsGluCl in BPH. (C) Egg production after GluCl gene knockdown in EB-treated and untreated BPH. (D) The JH III titer of BPH females after GluCl gene silencing as quantified using the ELISA method. (E-I) Expression patterns of selected JH-related genes (JHAMT, Met, Kr-h1 and JHE) in GluCl silenced BPH. (J) Schematic of the proposed role of GluCl as a molecular target of EB and EB-enhanced reproduction in BPH. The question mark indicates one or more possible additional signals. All data are presented as means ± s.e.m. Student’s t test was used to compare the two samples. One-way ANOVA with Tukey’s multiple comparisons test was used to compare more than two samples. ns, no significant difference; Asterisks indicate values significantly different from the control (ns, no significant; *p < 0.05 and **p < 0.01). Different lower-case letters above the bars indicate significant differences (p < 0.05).

Added in the revised version: Phylogenetic analysis of glutamate-gated chloride channels in different species. The numbers at the nodes of the branches represent the percentage bootstrap support (1000 replications) for each branch. The Sogatella furcifera GABA-gated chloride channel and Nilaparvata lugens nAchR were used as outgroup. Alignment was performed with amino acid sequences from TM1-7. The receptor names are listed in the tree. The accession numbers of the sequences used for this phylogenetic tree are listed in Figure 5-figure supplement 1 source data.

Added in the revised version: The expression of AT, AstA, AstB, AstCC, AstCCC, AstAR, AstBR and AstCR in BPH injected with dsGluCl or dsgfp. All data are presented as means ± s.e.m. Student’s t test was used to compare the two samples. ns, no significant. Different lower-case letters above the bars indicate significant differences (p < 0.05).

Added in the revised version: The expression of GluCl after BPH female adult treated with EB. All data are presented as means ± s.e.m. Student’s t test was used to compare the two samples, *, p < 0.05.

Discussion

Earlier version: Pesticide-induced resurgence of pest insects is a serious problem in rice and several other cropping systems (3). However, the mechanisms underpinning pesticide-enhanced reproduction in insects remain poorly understood. Here we reveal that a suite of molecular actors underlie this trait that, in combination, mediate profound physiological changes in the reproductive fitness of female BPH. Our data provide fundamental insights into the molecular mechanisms by which xenobiotics modify insect reproduction and have applied implications for the control of a highly damaging crop pest. We discuss these topics below.

Revised version: Pesticide-induced resurgence of pest insects is a serious problem in rice and several other crops (Wu et al., 2020). However, the mechanisms underpinning pesticide-enhanced reproduction in insects remain poorly understood. Here we reveal that a suite of molecular actors underlie this trait that, in combination, mediate profound physiological changes in the reproductive fitness of female BPHs. Our data provide fundamental insights into the molecular mechanisms by which xenobiotics modify insect reproduction and have applied implications for the control of a highly damaging crop pest. We discuss these topics below.

Earlier version: Sublethal doses of GluCl modulators, EB and abamectin, stimulates fecundity in BPH

Revised version: Sublethal doses of the GluCl modulators, EB and abamectin, stimulates fecundity in BPH

Earlier version: We show that in both contact and systemic assays EB and abamectin stimulate reproduction in BPH. Thus, insecticide-enhanced reproduction is likely a key factor in the BPH resurgence observed when farmers use EB and abamectin to control leaffolders in rice crops in China (29). Although this is the first report of sublethal doses of avermectins enhancing insect fecundity, our findings are consistent with previous studies which have shown that certain insecticides, herbicides and fungicides stimulate BPH reproduction (3, 7, 14, 63–68). Intriguingly, we show that EB only induces fecundity in female adults and is specific to BPH, with EB exposure failing to enhance reproduction in two related species, the small brown planthopper, L. striatellus and the white backed planthopper, S. furcifera, or the model insect D. melanogaster. Thus, the mechanisms underpinning this trait appear to act exclusively on female BPH and may be specific to this species.

Revised version: We show that in both contact and systemic assays the avermectins EB and abamectin stimulate reproduction in BPHs. Thus, insecticide-enhanced reproduction is likely a key factor in the BPH resurgence observed when farmers use EB and abamectin to control leaffolders in rice crops in China (Yao et al., 2017). Although this is the first report of sublethal doses of avermectins enhancing insect fecundity, our findings are consistent with previous studies, which have shown that certain insecticides, herbicides and fungicides stimulate BPH reproduction (Cheng et al., 2014; Ge et al., 2011; Ge et al., 2013; Wan et al., 2013; Wang et al., 2010; Wu et al., 2020; Yang et al., 2019; Zhang et al., 2014; Zhao et al., 2011). Intriguingly, we show that EB only induces fecundity in female adults and is specific to BPH, with EB exposure failing to enhance reproduction in two related species, the small brown planthopper, L. striatellus and the white backed planthopper, S. furcifera, or the model insect D. melanogaster. Thus, the mechanisms underpinning this trait appear to act exclusively on female BPHs and may be specific to this species. Nevertheless, our findings are of general interest since they link the molecular EB target, the GluCl channel and juvenile hormone signaling to fecundity.

Earlier version: Pesticides may stimulate insect reproduction through a variety of physiological and molecular mechanisms. Our data reveal that exposure to sub-lethal concentrations of EB results in profound changes to female BPH fitness, leading to increases in female weight, total protein content, cholesterol content, sugar content and egg production and decreases in duration of the preoviposition period. Some of these findings exhibit parallels with previous studies, which demonstrated that treating third-instar BPH nymphs with either deltamethrin, triazophos, or imidacloprid led to increased soluble sugar levels in the corresponding adults (69). Such metabolites provide the energy that drives BPH reproduction and resurgence (3). Thus, together with prior work, our results suggest that pesticides associated with resurgence stimulate nutritional intake in BPH to fuel enhanced energy-intensive reproduction.

Revised version: Pesticides may stimulate insect reproduction through a variety of physiological and molecular mechanisms. Our data reveal that exposure to sub-lethal concentrations of EB results in profound changes to female BPH fitness, leading to increases in female weight, contents of total protein, cholesterol, and sugar, as well as egg production and decreases in duration of the preoviposition period. Some of these findings exhibit parallels with previous studies, which demonstrated that treating third-instar BPH nymphs with either deltamethrin, triazophos, or imidacloprid led to increased soluble sugar levels in the corresponding adults (Yin et al., 2014). Such metabolites provide the energy that drives BPH reproduction and hence resurgence (Wu et al., 2020). Thus, together with prior work, our results suggest that pesticides associated with resurgence stimulate nutritional intake in BPH to fuel enhanced energy-intensive reproduction.

The JH signaling pathway plays an essential role in EB-induced fecundity in BPH

Earlier version: JH is a pleiotropic hormone which plays important roles in development and reproduction in insects (19, 70). Circulating JH titers are regulated by factors that control JH production in the corpora allata including biosynthetic enzymes and catabolic enzymes that regulate JH levels. Our results show that EB increases circulating JH III titers in BPH females over 2–4 days after eclosion (DAE) and promotes ovarian development. Previous studies have reported that triazophos and deltamethrin treatments also lead to increased circulating JH III titers in BPH females over 1–3 days post emergence. Similarly, jinggangmycin treatments were found to lead to increased JH titers (by approximately 45–50%) in BPH females over two days post emergence (71). Thus our findings, in combination with these previous studies, demonstrate that insecticide treatments can have dramatic effects on the regulation of key insect hormones involved in pest reproduction which can in turn drive pesticide resurgence.

Revised version: JH is a pleiotropic hormone which plays important roles in development and reproduction in insects (Riddiford, 2008; Santos et al., 2019). Circulating JH titers are regulated by factors that control JH production in the corpora allata including biosynthetic enzymes and catabolic enzymes that regulate JH levels. Our results show that EB increases circulating JH III titers in BPH females over 2–4 days after eclosion (DAE) and promotes ovarian development. Previous studies have reported that triazophos and deltamethrin treatments also lead to increased circulating JH III titers in BPH females over 1–3 days post emergence. Similarly, jinggangmycin treatments were found to lead to increased JH titers (by approximately 45–50%) in BPH females over two days post emergence (Xu et al., 2016). Thus, our findings, in combination with these previous studies, demonstrate that insecticide treatments can have dramatic effects on the regulation of key insect hormones involved in pest reproduction, which can in turn drive pesticide resurgence.

Earlier version: Although increased JH titers following pesticide exposure have been correlated with reduced levels of active JH esterase during the first three days PE (6), the type and number of mechanisms mediating the observed increase in hormone titer has remained an open question. Our data reveal that elevated JH titer in EB-exposed BPH is associated with the upregulation of genes that encode biosynthetic enzymes for JH (JHAMT) and downstream signaling genes that can induce vg gene expression (met and kr-h1). Using RNAi we provide functional evidence of the role of these genes in the regulation of JH III and fecundity of female BPH, and demonstrate that EB can restore the reduction in egg production resulting from the knockdown of JHAMT, met and kr-h1 expression.

Revised version: Although increased JH titers following pesticide exposure have been correlated with reduced levels of active JH esterase during the first three days PE (Ge et al., 2010), the type and number of mechanisms mediating the observed increase in hormone titer has remained an open question. Our data reveal that elevated JH titer in EB-exposed BPH is associated with the upregulation of genes that encode biosynthetic enzymes for JH (JHAMT) and downstream signaling genes that can induce vg gene expression (kr-h1). Using RNAi we provide functional evidence of the role of these genes in the regulation of JH III and fecundity of female BPHs, and demonstrate that EB can restore the reduction in egg production resulting from the knockdown of JHAMT, met and kr-h1 expression.

Earlier version: JHMAT is an enzyme that catalyzes the conversion of inactive precursors of JH to active JH in the final stages of JH biosynthesis (36, 38). Interestingly, while it has not been previously implicated in pesticide resurgence, treatment of the stored product pest Sitotroga cerealella with diallyl trisulfide, an insecticidal compound in garlic essential oil, was found to increase JHAMT mRNA levels (72). Because JHMAT is the key rate-limiting enzyme in regulation of JH titer our results suggest that its enhanced expression is a key molecular mechanism of pesticide resurgence in BPH.

Revised version: JHMAT is an enzyme that catalyzes the conversion of inactive precursors of JH to active JH in the final stages of JH biosynthesis (Nouzova et al., 2018; Shinoda and Itoyama, 2003). Interestingly, while it has not been previously implicated in pesticide resurgence, treatment of the stored product pest Sitotroga cerealella with diallyl trisulfide, an insecticidal compound in garlic essential oil, was found to increase JHAMT mRNA levels (Shah et al., 2022). Because JHMAT is the key rate-limiting enzyme in regulation of JH titer our results suggest that its enhanced expression is a key molecular mechanism of pesticide resurgence in BPH.

Earlier version: Met is a ligand-activated member of the basic helix–loop–helix Per/Arnt/Sim (bHLH-PAS) transcription factors and is the intracellular receptor for JH (39, 40). Kr-h1 is a zinc finger protein that acts downstream of Met and is expressed in response to JH signaling. Although the genes encoding these proteins have not been previously linked to pesticide resurgence, our finding that they are upregulated following EB exposure, and demonstration of their role of in promoting fecundity, is consistent with previous studies. Specifically, treatment of BPH with JH III or the insecticidal analogues methoprene or pyriproxifen was found to induce the expression of Kr-h1 (73). Furthermore, knockdown of Met and Kr-h1 in BPH brachypterous females was found to result in delayed ovariole development and this was significantly more pronounced than the response observed in BPH treated separately with dsNlMet or dsNlKr-h1 (74). This finding provides evidence of a possible interaction between Met and Kr-h1 and, in combination with our data, suggests that Met and Kr-h1 may act in concert to mediate EB-enhanced fecundity.

Revised version: Met is a ligand-activated member of the basic helix–loop–helix Per/Arnt/Sim (bHLH-PAS) transcription factors and is the intracellular receptor for JH (Riddiford, 2020; Zhu et al., 2019). Kr-h1 is a zinc finger protein that acts downstream of Met and is expressed in response to JH signaling. Although the genes encoding these proteins have not been previously linked to pesticide resurgence, our finding that they are upregulated following EB exposure, and demonstration of their role of in promoting fecundity, is consistent with previous studies (Santos et al., 2019). Specifically, treatment of BPH with JH III or the insecticidal analogs methoprene or pyriproxifen was found to induce the expression of Kr-h1 (Jin et al., 2014). Furthermore, knockdown of Met and Kr-h1 in BPH brachypterous females was found to result in delayed ovariole maturation and this was significantly more pronounced than the response observed in BPH treated separately with dsNlMet or dsNlKr-h1 (Lin et al., 2015). This finding provides evidence of a possible interaction between Met and Kr-h1 and, in combination with our data, suggests that Met and Kr-h1 may act in concert to mediate EB-enhanced fecundity. However, our results showed that EB treatment can mildly increase (about 2-fold) expression of the Met gene in brown planthoppers (Figure 3G). Our data furthermore indicated that Met and FAMeT expression levels were not drastically influenced by EB compared with kr-h1 and vg (Figure 3H and I). It should be mentioned that JH action does not directly result in the increase of Met. However, we cannot rule out the possibility of other factors (indirect effects), induced by EB treatment, increases the mRNA expression level of Met. One recent paper reported that downregulation of transcription factor CncC will increase met expression in beetles (see Figure 6A in the reference) (Jiang et al., 2023). And many studies have reported that insecticide treatment activates the CncC gene signaling pathway, which regulates detoxification gene expression (Amezian et al., 2023; Fu et al., 2024; Hu et al., 2021). Hence, it is possible that EB might influence the CncC gene pathway, which in turn induces met expression. This EB effect on met upregulation may be similar to the upregulation of GluCl and some other secondary effects (see below).

Added in the revised version: Schematic of the proposed regulatory pathway of EB-enhanced fecundity in BPH.

Added in the revised version: Emamectin benzoate (EB) exposure results in the upregulation of genes that promote JH signaling pathway (JHAMT and Kr-h1) and the downregulation of genes that inhibit it (allatostatin, AstA and allatostatin A receptor, AstAR). This transcriptome reprograming is dependent on the action of EB on its molecular target the glutamate-gated chloride channel (GluCl) receptor. The resulting increased JH titer promotes vg synthesis and increased fecundity in EB exposed insects. We observe significant cross-talk in the expression of genes that inhibit JH production and those that promote it, with AstAR inhibiting the expression of JHAMT, Met and Kr-h1 and GluCl activating the expression of JHAMT which is responsible for JH synthesis, and the JH signalling downstream genes Met and Kr-h1.

Added in the revised version: It must be mentioned that although we exposed the BPH to EB at 4th instar, we make the insect feed on the EB-treated rice plants for four days. After that, the insect will develop into 5th instar, the final nymphal stage of brown planthopper. Since BPH do not have a pupal stage, this might cause the EB presented to the insects last a longer time even in the adult stage. Besides this, we found that EB treatment will increase the weight of adult females, which indicates that EB might increase food intake in BPHs that might produce more insulin peptide. Hence, insulin might increase the JH synthesis at the adult stage.

Earlier version: EB-induced fecundity in BPH is dependent on the allatostatin signaling pathway

Revised version: EB-induced fecundity in BPHs is dependent on the allatostatin A signaling pathway

Earlier version: In addition to regulatory proteins that promote JH production, insects have peptides that inhibit JH biosynthesis. These include the allatostatins: FGLamides (FGLa; AstA), the W(X)6Wamides (AstB), and the PISCFs (AstC) (50, 51, 57, 58, 75, 76). Interestingly, our results showed that EB exposure results in the marked downregulation of the expression of the genes encoding allatostatin AstA and its receptor AstAR. We provide clear evidence of the functional impact of this on JH synthesis and BPH fecundity by: i) demonstrating that RNAi knockdown of AstAR expression results in increased JH titer and enhanced female egg production, and, ii) showing that injection of female BPH with synthetic AStA peptide reduces JH titer and decreases egg production. Thus, our data provide unequivocal evidence that AstA is a key inhibitor of JH production in BPH. This finding is consistent with previous work which has shown that FGLa/ASTs inhibit JH biosynthesis in cockroaches, and termites (46, 77). To our knowledge, our study is the first report of insecticides inhibiting the expression of the neuropeptide receptor, AstAR, and linking this to increases in JH titer and enhanced reproduction in insects.

Revised version: In addition to regulatory proteins that promote JH production, insects utilize peptides that inhibit JH biosynthesis. These include the allatostatins: FGLamides (FGLa; AstA), the W(X)6Wamides (AstB), and the PISCFs (AstC) (Nässel, 2002; Stay et al., 1996; Stay and Tobe, 2007; Tanaka et al., 2014; Wegener and Chen, 2022; Zhang et al., 2021). Interestingly, our results showed that EB exposure results in a marked downregulation of the expression of the genes encoding allatostatin AstA and its receptor AstAR. We provide evidence for the functional impact of this on JH synthesis and BPH fecundity by: i) demonstrating that RNAi knockdown of AstAR expression results in increased JH titer and enhanced egg production, and, ii) showing that injection of female BPH with synthetic AstA peptide reduces JH titer and decreases egg production. Thus, our data provide unequivocal evidence that AstA is a key inhibitor of JH production in BPH. This finding is consistent with previous work which has shown that FGLa/ASTs (AstA) inhibit JH biosynthesis in cockroaches, and termites (Woodhead et al., 1989; Yagi et al., 2005). To our knowledge, our study is the first report of insecticides down-regulating the expression of the neuropeptide receptor, AstAR, and linking this to increases in JH titer and enhanced reproduction in insects.

Earlier version: Interestingly knockdown of AstAR resulted in significant increases in the expression of genes involved in JH synthesis/signaling including FAMeT and JHAMT Met and Kr-h1. Related to this finding, previous work has shown that knockdown of the AstA receptor gene, Dar-2, in D. melanogaster results in changes in the expression of genes encoding Drosophila insulin-like peptides (DILPs) and adipokinetic hormone (AKH) signaling proteins (78). Together with our findings, this demonstrates that AstA receptors may modulate the expression of numerous downstream genes involves in metabolism, energy store and reproduction. In the case of pesticide resurgence our results imply significant cross-talk in the expression of genes that inhibit JH production and those that promote it.

Revised version: Interestingly knockdown of AstAR resulted in significant increases in the expression of genes involved in JH synthesis/signaling including JHAMT, Met and Kr-h1. Related to this finding, previous work has shown that knockdown of the AstA receptor gene, Dar-2, in D. melanogaster results in changes in the expression of genes encoding Drosophila insulin-like peptides (DILPs) and adipokinetic hormone (AKH) (Hentze et al., 2015). Together with our findings, this demonstrates that AstA receptors may modulate the expression of numerous downstream genes involves in metabolism, energy store and reproduction. In the case of pesticide resurgence our results imply significant cross-talk in the expression of genes that inhibit JH production and those that promote it.

The GluCl plays an essential role in EB-induced fecundity in BPH

Added in the revised version: EB and abamectin are allosteric modulators of GluCls (Sparks et al., 2020; Sparks et al., 2021; Wu et al., 2017). Our data revealed that EB exposure increases the expression of the GluCl in adult stage of BPH, and knockdown of GluCl expression resulted in a reduction in both JH levels and egg production and egg laying. There are several studies showing that insects treated with insecticides display increases in the expression of target genes. For example, the relative expression level of the ryanodine receptor gene of the rice stem borer, Chilo suppressalis was increased 10-fold after treatment with chlorantraniliprole, an insecticide which targets the ryanodine receptor (Peng et al., 2017). Besides this, in Drosophila, starvation (and low insulin) elevates the transcription level of the short neuropeptide F and tachykinin receptors (Ko et al., 2015; Root et al., 2011). In BPHs, reduction in mRNA and protein expression of a nicotinic acetylcholine receptor α8 subunit is associated with resistance to imidacloprid (Zhang et al., 2015). RNA interference knockdown of the α8 gene decreased the sensitivity of N. lugens to imidacloprid (Zhang et al., 2015). Hence, expression of receptor genes may be regulated by diverse factors, including insecticide treatment. In our case, we found that EB upregulates its target gene GluCl. However, we do not suggest that EB functions as transcriptional activator of GluCl and we have no data to suggest a mechanism for how EB treatment changes the expression of GluCl in the BPH. Considering that EB exposure in our experiments is lasting several days, it might be an indirect (or secondary) effect caused by other factors downstream GluCl that affect channel expression Possibly the allosteric interaction with GluCl by EB renders the channel dysfunctional and the cellular response is to upregulate the channel/receptor to compensate.

Earlier version: EB and abamectin are allosteric modulators of GluCls (60–62). Our data revealed that EB exposure modifies expression of the GluCl in BPH, and knockdown of GluCl expression resulted in a reduction in both JH levels and egg production. Interestingly, the GluCl has been reported to inhibit the biosynthesis of JH in the cockroach, Diploptera punctata (28). Recent work has also reported that modulation of glutamatergic signals may contribute to the photoperiodic control of reproduction in bean bug, Riptortus pedestris (79). Interestingly, work on D. punctata has revealed that application of the GluCl channel agonist ivermectin, which like EB belongs to the avermectin family, caused a decline in JH synthesis in corpus allatum glands (28). While the inhibitory effect of ivermectin observed in this previous study differs from the activating effect of EB we observed in our study, it is consistent with our finding of a role for GluCl channel in the regulation of JH regulation. Interestingly, we found that knockdown of GluCl gene expression results in the down-regulation of JHAMT, Met and Kr-h1, further revealing significant convergent relationships between genes underpinning pesticide resurgence.

Revised version: Interestingly, the GluCl has been reported to inhibit the biosynthesis of JH in the cockroach, Diploptera punctata (Liu et al., 2005). Recent work has also reported that modulation of glutamatergic signals may contribute to the photoperiodic control of reproduction in bean bug, Riptortus pedestris (Hasebe and Shiga, 2022). Furthermore, work on D. punctata has revealed that application of the GluCl channel agonist ivermectin, which like EB belongs to the avermectin family, caused a decline in JH synthesis in the corpus allatum (Liu et al., 2005). While the inhibitory effect of ivermectin observed in this previous study differs from the activating effect of EB we observed in our study, it is consistent with our finding of a role for the GluCl channel in the regulation of JH regulation. Interestingly, we found that knockdown of the GluCl gene expression results in the down-regulation of JHAMT, Met and Kr-h1, further revealing significant convergent relationships between genes underpinning pesticide resurgence.

Conclusion

Earlier version: Our study has revealed a diverse suite of genes that act in combination to enhance JH titer and thus fecundity following BPH exposure to EB. A schematic of how these factors promote ovarian development in the adult stage of N. lugens through the JH signaling pathway is provided in Figure 6. Our findings provide the foundation for further work to understand exactly how these genes interact and the mechanisms by which their expression is activated or repressed by EB. Furthermore, our findings provide fundamental insights into the molecular response in insects to xenobiotic stress and illustrate that pesticides can have unexpected and negative impacts on pest populations. In this regard our findings also have applied implications for the control of a highly damaging crop pest. Previous studies have reported that avermectins such as abamectin are toxic to the wolf spider Pardosa pseudoannulata, which is the main predator of BPH in rice crops (80, 81). Thus, these insecticides both stimulate reproduction in BPH while killing their natural enemies providing a ‘perfect storm’ for damaging BPH outbreaks. Based on these findings, to avoid BPH resurgence, we suggest that the GluCls, EB and abamectin, should not be (or rarely be) applied to rice plants at growth stages when BPHs are present. On a more optimistic note, our findings have identified numerous genes that play key roles in BPH reproduction and thus represent promising targets for the development of novel controls against this important pest.

Revised version: Our study has revealed that exposure of BPHs to EB affects a diverse suite of genes that act in combination to enhance JH titer and thus increase female fecundity. A schematic of how these factors promote ovarian maturation in the adult stage of N. lugens through the AstA and JH signaling pathway is provided in Figure 6. Our findings provide the foundation for further work to understand how these genes interact and the mechanisms by which their expression is activated or repressed by EB. Furthermore, our findings provide fundamental insights into the molecular response in insects to xenobiotic stress and illustrate that pesticides can have unexpected and negative impacts on pest populations. In this regard our findings also have applied implications for the control of a highly damaging crop pest. Previous studies have reported that avermectins such as abamectin are toxic to the wolf spider Pardosa pseudoannulata, which is the main predator of BPH in rice crops (Chen et al., 2017; Sogawa, 2015). Thus, these insecticides both stimulate reproduction in BPH while killing their natural enemies providing a ‘perfect storm’ for inducing damaging BPH outbreaks. Based on these findings, to avoid BPH resurgence, we suggest that the agents, EB and abamectin, should not be (or rarely be) applied to rice plants at growth stages when BPHs are present. On a more optimistic note, our findings have identified numerous genes that play key roles in BPH reproduction and thus represent promising targets for the development of novel controls against this important pest.

Removed in the revised version: Schematic of the proposed regulatory pathway of EB-enhanced fecundity in BPH.

Removed in the revised version: Emamectin benzoate (EB) exposure results in the upregulation of genes that promote JH production (JHAMT, Met and Kr-h1) and the downregulation of genes that inhibit it (allatostatin, AstA and allatostatin A receptor, AstAR). This transcriptome reprograming is dependent on the action of EB on its molecular target the glutamate-gated chloride channel (GluCl) receptor. The resulting increased JH titer promotes vg synthesis and increased fecundity in EB exposed insects. We observe significant cross-talk in the expression of genes that inhibit JH production and those that promote it, with AstAR inhibiting the expression of JHAMT, Met and Kr-h1 and GluCl activating the expression of JHAMT which is responsible for JH synthesis, and the JH signalling downstream genes Met and Kr-h1.

Materials and methods

Insects

BPH was initially collected from Wanan, JiangXi Province in 2020, reared on ‘Taichung Native 1’ (TN1) rice seedlings in the laboratory without exposure to any insecticides. The strain was maintained in a climatic chamber at 27 ± 1°C, with relative humidity of 70 ± 10% and a light: dark = 16 h: 8 h photoperiod.

Chemicals

Emamectin benzoate (95.2%) was obtained from Hebei Weiyuan Co., Ltd. (Hebei, China). Abamectin (96.8%) was obtained from Hebei Weiyuan Co., Ltd. (Hebei, China). Pyriproxyfen (98%) was obtained from Shanghai Shengnong Co., Ltd. (Shanghai, China). Methoprene (S)-(+) (mx7457-100mg) was purchased from Shanghai MaoKang Biotechnology Co., Ltd., (Shanghai, China). Juvenile hormone standard sample (J912305-10mg) was purchased from Shanghai Macklin Biotechnology Co., Ltd., (Shanghai, China).

Bioassay

Different life stages of insects were used to perform bioassay to investigate the effects of insecticide on nymphs and adults. To test whether treatment of the nymph stage of insects would promote reproduction in female, we used 4th instar nymphs of BPH or 3rd instar nymph of Laodelphax striatellus and Sogatella furcifera to perform bioassays. To test whether treatment of the adult stage of insects would promote reproduction in female, we used newly emerged male and female BPH.

Systemic route

The rice-seeding dipping bioassay method was used to evaluate the susceptibility of BPH, L. striatellus and S. furcifera to EB. Technical-grade insecticides were dissolved in acetone as stock solution then diluted in a series of six concentrations with water containing 0.1% Triton. Selected rice seedlings at the 6-8 cm growth stage were dipped in insecticide solutions for 30 s and then air-dried at room temperature. The roots of the rice seedlings were wrapped with cotton strips and placed seedlings placed in a plastic cup 5 cm in diameter. Fifteen insects were introduced into each plastic cup for each replicate. The top of the cup was then selaed with gauze to prevent escape. All experiments comprised at least three biological replicates. Control rice seedings were treated with 0.1% Triton X-100 water solution only. All treatments were maintained under standard conditions of 27 ± 1 °C and 70–80% relative humidity with a 16 h light/8 h dark photoperiod. Mortality was assessed after 4 d for N. lugens or 2 d for L. striatellus and S. furcifera after treatment with insecticides. The insects were considered dead if they were unable to move after a gentle prodding with a fine brush.

Earlier version: For Drosophila larvae bioassay, we adopted a method described previously in our lab with minor modifications (82). Briefly, twenty third instar larvae were placed in fly vials containing fly food (based on corn powder, brown sugar, yeast and agar) supplemented with EB of different concentrations. Four concentrations (LC10, LC30 and LC50) were tested together with a negative (no insecticide) control. For Drosophila adult bioassays, we selected virgin females three days after eclosion. Several concentrations were overlaid onto fly food in standard Drosophila vials and allowed to dry overnight at room temperature. 15 adult flies (three days after eclosion) were then added to each vial and mortality assessed after 2 d. Four replicates were carried out for each concentration. Control mortality was assessed using vials containing food with solvent minus insecticide.

Revised version: For Drosophila larvae bioassay, we adopted a method described previously in our lab with minor modifications (Huang et al., 2020). Briefly, twenty third instar larvae were placed in fly vials containing fly food (based on corn powder, brown sugar, yeast and agar) supplemented with EB of different concentrations. Four concentrations (LC10, LC30 and LC50) were tested together with a negative (no insecticide) control. For Drosophila adult bioassays, we selected virgin females three days after eclosion. Several concentrations were overlaid onto fly food in standard Drosophila vials and allowed to dry overnight at room temperature. 15 adult flies (three days after eclosion) were then added to each vial and mortality assessed after 2 d. Four replicates were carried out for each concentration. Control mortality was assessed using vials containing food with solvent minus insecticide.

Contact route

Earlier version: For topical bioassays working insecticide solutions were prepared in acetone. 4th instar nymphs or newly emerged males/females were anesthetized with carbon dioxide for 5 s, and then 0.04 μl/insect test solution applied topically to the dorsal plates with a capillary micro-applicator (29). Insects were then placed in an artificial climate incubator with a temperature of 27±1°C, a photoperiod of 16:8 h (L:D), and a humidity of 70%±10%. Mortality was determined 2 d after treatment. Data with over 20% mortality in the control treatment were discarded, and the assay was repeated at least three times.

Revised version: For topical bioassays working insecticide solutions were prepared in acetone. 4th instar nymphs or newly emerged males/females were anesthetized with carbon dioxide for 5 s, and then 0.04 μl/insect test solution applied topically to the dorsal plates with a capillary micro-applicator (Yao et al., 2017). Insects were then placed in an artificial climate incubator with a temperature of 27±1°C, a photoperiod of 16:8 h (L:D), and a humidity of 70%±10%. Mortality was determined 2 d after treatment. Data with over 20% mortality in the control treatment were discarded, and the assay was repeated at least three times.

Fecundity assays

Fourth instar nymphs of BPH and 3rd instar nymph of L. striatellus and S. furcifera were treated with the LC15 and LC50 of EB, and then transferred to fresh rice seedlings. After eclosion, the adults were used in the following experiment. Newly emerged treated adults and untreated adults were paired to produce four groups: untreated males and untreated females (♂ck×♀ck; ck indicates untreated); untreated males and treated females (♂ck×♀t; t indicates insecticide treated); treated males and untreated females (♂t×♀ck); treated males and treated females (♂t×♀t). Each group comprised at least 10 mating pairs. All pairs were transferred to glass cylinders (diameter 5 cm and height 20 cm) containing rice plants (25 days old almost 20 cm high) as a food source for eleven days. The number of eggs and nymphs in plants were counted by a microscope.

Earlier version: For Drosophila egg-laying assay, we adopted our previous method (83). Briefly, insecticide-treated virgin females were paired with untreated males for three days and then the mated females transferred into the Drosophila ovipositing apparatus. The number of eggs were counted after 16 hours.

Revised version: For Drosophila egg-laying assay, we adopted our previous method (Wu et al., 2019). Briefly, insecticide-treated virgin females were paired with untreated males for three days and then the mated females transferred into the Drosophila ovipositing apparatus. The number of eggs were counted after 16 hours.

Fitness analysis

Earlier version: The fitness of EB-treated BPH were analyzed using methods reported previously (10). We selected two groups, (♂ck×♀ck) and (♂ck×♀t), to study the effects of the LC50 concentration of EB on BPH fitness. In the case of systemic exposure, 4th instar nymphs of BPH were treated with the LC50 of EB for 4 days and then transferred to tubes containing untreated rice plants for individual rearing. The rice plants were replaced every three days with untreated plants. The emergence ratio, female ratio, preoviposition period, female longevity, brachypterism female ratio and female weight were calculated.

Revised version: The fitness of EB-treated BPH were analyzed using methods reported previously (Zeng et al., 2023). We selected two groups, (♂ck×♀ck) and (♂ck×♀t), to study the effects of the LC50 concentration of EB on BPH fitness. In the case of systemic exposure, 4th instar nymphs of BPH were treated with the LC50 of EB for 4 days and then transferred to tubes containing untreated rice plants for individual rearing. The rice plants were replaced every three days with untreated plants. The emergence ratio, female ratio, preoviposition period, female longevity, brachypterism female ratio and female weight were calculated.

Examination of ovarian development

Earlier version: Adult females from ♂ck×♀ck control and ♂ck×♀t group on 1, 3, 5, 7 DAE were dissected for observe the ovarian development. The detained eggs in ovary were photographed and recorded. Each group has at least fifteen replicates.

Revised version: Adult females from ♂ck×♀ck control and ♂ck×♀t group on 1, 3, 5, 7 DAE were dissected for observe the ovarian development. The mature eggs in ovary were photographed and recorded. Each group has at least fifteen replicates.

To examination whether EB treated impairs egg maturation, we dissected untreated or EB-treated ovaries and fixed them in 4% paraformaldehyde (PFA) in phosphate buffered saline (PBS) for 30 min at room temperature. After four washes of 10 min (4 × 10 min) in PAT3 (PBS with 0.5% Triton X-100 and 0.5% bovine serum albumin), the ovaries were then incubated with DAPI (4’,6-diamidino-2-phenylindole, 100 nM) and Actin-stain 670 Fluorescent Phalloidin (200 nM). Imaging was performed using Zeiss LSM980 confocal laser microscope.

Measurements of glycogen, triglyceride, total protein content, cholesterol and four sugars

The content of glycogen, triglyceride, cholesterol and total protein was determined by spectrophotometry at 620 nm, 510 nm, 510 nm and 562 nm respectively using the glycogen assay kit (A043-1-1), triglyceride reagent kit (A110-1-1), cholesterol assay kit (A111-1-1) and total protein assay kit (A045-2-2) obtained from Nanjing Jiancheng Bioengineering Institute following to the manufacturer’s instructions. The determined results were normalized to the protein content in the sample, which was determined using BCA Protein Assay Reagent Kit (Thermo Scientific, Waltham, USA). Each sample contained tissue extracts from five adult female BPH, with three biological replicates per sample.

To assess the content of four sugars (sucrose, glucose, fructose and trehalose) in the extract of BPH tissue, the same extraction method was used as above. Sugar content was quantified using the colorimetric method by scanning at 290 nm, 505 nm, 285 nm and 620 nm respectively using the sucrose reagent kit (A099-1-1), glucose reagent kit (F006-1-1), fructose reagent kit (A085-1-1) and trehalose reagent kit (A150-1-1) obtained from Nanjing Jiancheng Bioengineering Institute based on the manufacturer’s instructions. Each sample contained tissue extracts from five adult female N. lugens, with three biological replicates per sample.

Determination of Juvenile hormone III and ecdysone titers of BPH by ELISA

The titer of Juvenile hormone III in BPH was measured using the Juvenile hormone ELISA Kit (Lot Number: HLE92086, Shanghai HaLing Biotechnology Co., Ltd., Shanghai, China) which employs the dual-antibody sandwich ELISA method. The titer of ecdysone in BPH were measured using the ecdysone ELISA Kit (Lot Number: ZK14705, Shanghai ZhenKe Biotechnology Co., Ltd., Shanghai, China). At least three biological replicates were employed for each treatment.

Determination of Juvenile hormone III titer in BPH using HPLC-MS/MS

Earlier version: The whole bodies of 5 individuals BPH were mixed with 1 ml of n-hexane, followed by centrifugation at 10,000×g for 10 min, the upper hexane layer was then dried with nitrogen, dissolved in methanol and sonicated for 10 min, after centrifugation at 10000×g for 10 min, the supernatant was collected through the organic filter membrane of 0.22 μm into 2 ml vials for JH Ⅲ determination. JH Ⅲ standard sample (J912305-10 mg) purchased from (Shanghai McLean Biochemical Technology Co. Ltd), dissolved in methyl alcohol as stock solution 10,000 mg/L was diluted in a series of six concentration gradients to serve as a reference. Liquid chromatography-tandem mass spectrometry (LC-MS/MS) was then carried out using UPLC Xevo TQ-S Micro (Waters technology), quantitative method according to the external standard, the chromatographic column was EC-C18 (4.6 mm×150 mm, 2.7 μm), column temperature was 30°C, injection volume was 20 μl, elution flow rate was 0.3 ml/min, and the mobile phase was acetonitrile:formic acid water (90:10), detection wavelength was 218 nm, the peak height was used for quantification.

Revised version: The whole bodies of 5 individuals BPH were mixed with 1 ml of n-hexane, followed by centrifugation at 10,000×g for 10 min, the upper hexane layer was then dried with nitrogen, dissolved in methanol and sonicated for 10 min, after centrifugation at 10000×g for 10 min, the supernatant was collected through the organic filter membrane of 0.22 μm into 2 ml vials for JH LJ determination. JH LJ standard sample (J912305-10 mg) purchased from (Shanghai McLean Biochemical Technology Co. Ltd), dissolved in methyl alcohol as stock solution 10,000 mg/L was diluted in a series of six concentration gradients to serve as a reference. Liquid chromatography-tandem mass spectrometry (LC-MS/MS) was then carried out using UPLC Xevo TQ-S Micro (Waters technology), quantitative method according to the external standard, the chromatographic column was EC-C18 (4.6 mm×150 mm, 2.7 μm), column temperature was 30°C, injection volume was 20 μl, elution flow rate was 0.3 ml/min, and the mobile phase was acetonitrile:formic acid water (90:10), detection wavelength was 218 nm, the peak height was used for quantification.

Cloning, sequence and phylogenetic analysis

The NCBI database and BLAST program were used to carry out sequence alignment and analysis. Open Reading Frames (ORFs) were predicted with EditSeq. Primers were designed using the primer design tool in NCBI. Total RNA Extraction was extracted from 30 adults BPH using TRIzol reagent (Invitrogen, Carlsbad, CA, USA) according to the manufacturer’s instructions. cDNA was synthesized using the Biotech M-MLV reverse transcription kit. Full-length gene sequences were amplified by PCR using cDNA as template and Rapid Taq Master Mix (Vazyme Biotech, Cat# P222-02). The PCR product was purified on a 1% agarose gel, cloned into pClone007 Simple Vector Kit (Tsingke Biotech, Cat# TSV-007S), and then sequenced using the 3730 XL DNA analyzer (Applied Biosystems, Carlsbad, CA, USA). Table S2 contains a list of the primers used in this study.

Earlier version: The exon and intron architectures of AT, AstA, AstB, AstCC and AstACCC were predicted based on the alignments of putative ORFs against their corresponding genomic sequences. Sequence similarity/annotations and orthologous gene searches were performed using BLAST programs available in NCBI. Multiple alignments of the complete amino acid sequences were performed with Clustal Omega (http://www.ebi.ac.uk/Tools/msa/clustalo). Phylogeny was conducted using the maximum likelihood technique to create phylogenetic trees and these were bootstrapped with 1000 replications were used using the MEGA 6 software (84).

Revised version: The exon and intron architectures of AT, AstA, AstB, AstCC and AstACCC were predicted based on the alignments of putative ORFs against their corresponding genomic sequences. Sequence similarity/annotations and orthologous gene searches were performed using BLAST programs available in NCBI. Multiple alignments of the complete amino acid sequences were performed with Clustal Omega (http://www.ebi.ac.uk/Tools/msa/clustalo). Phylogeny was conducted using the maximum likelihood technique to create phylogenetic trees and these were bootstrapped with 1000 replications were used using the MEGA 6 software (Tamura et al., 2013).

RNA interference

Double-stranded RNA (dsRNA) of gfp (green fluorescent protein), JHAMT (juvenile hormone acid O-methyltransferase,), Met (methoprene-tolerant), Kr-h1 (krüppel homolog 1), AstAR (allatostatin-A receptor) and GluCl (glutamate-gated chloride channel) was prepared using Ambion’s MEGAscript T7 kit instructions following the manufacturer’s instructions. The primer sequences for double-stranded RNA synthesis are listed in Table S2. Newly emerged females were injected with 40 nl (5,000 ng/μl) of double-stranded RNA of gfp (dsgfp) or double-stranded RNA of the target genes in the conjunctive part between prothorax and mesothorax of insects. In the RNAi experiments, BPH were then treated with the LC50 of EB 24h after dsRNA injection and the whole body sampled for qRT-PCR analysis.

Quantitative RT-PCR

Earlier version: Fourth instar nymphs of BPH were treated with EB after which total RNA was extracted from 5th instar nymphs and 1-7 day post-eclosion females of N. lugens using the methods detailed above. The HiScript® II RT SuperMix for qPCR (+gDNA wiper) kit from Vazyme, Nanjing, China, was used to create first-strand cDNA. Primer3 (http://bioinfo.ut.ee/primer3/) was used to design real-time quantitative PCR (qPCR) primers listed in Table S2. mRNA levels of candidate genes were detected by qPCR using the UltraSYBR Mixture (with ROX) Kit (CWBIO, Beijing, China). Each reaction contained 2 μL of cDNA template (500 ng), 1 μL each forward and reverse qPCR primer (10 μM), 10 μL of UltraSYBR mixture buffer, and 6 μL of RNase-free water. Q-PCR was run on an ABI 7500 Real-Time PCR System (Applied Biosystems, Foster City, CA) under the following conditions: 5 min at 95°C, followed by 40 cycles of 95°C for 10 s and 60°C for 30 s. Three independent biological replicates and four technical replicates were used in each qPCR experiment. The housekeeping genes 18S ribosomal RNA of BPH were selected to normalize the expression of candidate genes. The 2−ΔΔCt method (Ct represents the cycle threshold) was used to measure relative expression levels (85). Three biological replicates were used for statistical comparison between samples. Table S2 contains a list of the primers used in this study.

Revised version: Fourth instar nymphs of BPH were treated with EB after which total RNA was extracted from 5th instar nymphs and 1-7 day post-eclosion females of N. lugens using the methods detailed above. The HiScript® II RT SuperMix for qPCR (+gDNA wiper) kit from Vazyme, Nanjing, China, was used to create first-strand cDNA. Primer3 (http://bioinfo.ut.ee/primer3/) was used to design real-time quantitative PCR (qPCR) primers listed in Table S2. mRNA levels of candidate genes were detected by qPCR using the UltraSYBR Mixture (with ROX) Kit (CWBIO, Beijing, China). Each reaction contained 2 μL of cDNA template (500 ng), 1 μL each forward and reverse qPCR primer (10 μM), 10 μL of UltraSYBR mixture buffer, and 6 μL of RNase-free water. Q-PCR was run on an ABI 7500 Real-Time PCR System (Applied Biosystems, Foster City, CA) under the following conditions: 5 min at 95°C, followed by 40 cycles of 95°C for 10 s and 60°C for 30 s. Three independent biological replicates and four technical replicates were used in each qPCR experiment. The housekeeping genes 18S ribosomal RNA of BPH were selected to normalize the expression of candidate genes. The 2−ΔΔCt method (Ct represents the cycle threshold) was used to measure relative expression levels (Livak and Schmittgen, 2001). Three biological replicates were used for statistical comparison between samples. Table S2 contains a list of the primers used in this study.

Statistics

PoloPlus v2.0 (LeOra Sofware 2008) was used to calculate the lethal concentration (LC50) and 95% fiducial limits (95% F.L.). GraphPad Prism 8.0 software (GraphPad Software Inc., San Diego, USA) was used to generate graphs and perform statistical analysis of data. Data presented in this study were first verified for normal distribution using the D’Agostino–Pearson normality test. One-way analysis of variance (ANOVA) with Duncan’s multiple range test was used to test differences among multiple groups of normally distributed groups of data. Student’s t test was used to test the differences between two groups. If not normally distributed, Mann–Whitney test was used for pairwise comparisons, and Kruskal–Wallis test was used for comparisons among multiple groups, followed by Dunn’s multiple comparisons. All data are presented as mean ± s.e.m. The sample sizes and statistical tests used for each experiment are stated in the figures or figure legends.